Nuclear export modulates TDP-43 phase transition and cytoplasmic aggregation

Author response:

The following is the authors’ response to the original reviews.

Public Reviews:

Reviewer #1 (Public review):

In this paper, the authors use a doxycycline-inducible DLD1 cell line expressing a Clover-tagged RNA-binding-defective TDP-43 2KQ mutant that forms nuclear "anisosomes" (TDP-43 shell with HSP70 core) to carry out a small-molecule screen using the LOPAC 1280 library to identify compounds that reduce anisosome number or shift their morphology and dynamics. They also conducted a genome-wide siRNA screen to identify genetic modifiers of anisosome formation and dynamics. From these screens, the authors identify pathways in RNA splicing, translation, proteostasis (proteasome and HSP90), and nuclear transport, including XPO1. They then focus on XPO1 as their primary hit. Pharmacological inhibition of XPO1 using KPT-276, Verdinexor, and Leptomycin B reduces anisosome number while enlarging remaining condensates, which retain liquid-like behavior by FRAP and fusion assays. XPO1 overexpression causes fewer, enlarged TDP-43 puncta, including cytoplasmic puncta, with little or no FRAP recovery, interpreted as gel or solid-like aggregates. Anisosome induction reduces detectable nucleoplasmic XPO1 staining. Finally, the authors examine a homozygous TDP-43 K181E iPSC-derived forebrain organoid model, showing increased cytosolic pTDP-43 in K181E/K181E organoids compared to wild-type controls. Chronic low-dose KPT-276 reduces cytoplasmic pTDP-43 without changing total TDP-43 levels. Bulk RNA-seq shows only a modest fraction of dysregulated genes in K181E/K181E organoids are rescued by KPT-276. They conclude that nuclear export, via XPO1, is a key regulator of TDP-43 liquid-to-solid phase transitions and that cytoplasmic aggregation per se may contribute only modestly to TDP-43 proteinopathy, with RNA-processing defects being dominant.

We thank the reviewer for carefully summarizing our study.

The study presents well-executed chemical and genome-wide siRNA screens in a DLD1 TDP-43 2KQ anisosome model and follows up on nuclear transport, particularly XPO1, as a modulator of TDP-43 phase behavior and cytoplasmic aggregation. The screens are impressive in scale, and the microscopy and fluorescence recovery after photobleaching (FRAP) work is technically strong. However, the central mechanistic and disease-relevance claims are not yet sufficiently supported. There are major concerns about the heavy reliance on non-physiological, RNA-binding-defective, and acetylation-mimetic TDP-43 (2KQ) and a homozygous TDP-43 K181E organoid model. An underdeveloped and partly contradictory mechanistic link exists between XPO1 and TDP-43 phase transitions in the context of prior work showing TDP-43 is not a canonical XPO1 cargo. The paper also appears to overinterpret organoid data to conclude that cytoplasmic TDP-43 aggregation plays only a minor role in pathology, based largely on pTDP-43 antibody staining with limited sensitivity and relatively modest rescue readouts. A deeper mechanistic analysis and additional, more physiological validation are needed for this to reach the level of rigor and impact implied by the title and abstract. The work feels screen-rich but conceptually underdeveloped, with key claims outpacing the data. A major revision with substantial new data and tempering of conclusions is warranted. I outline several problematic areas below:

(1) The central mechanistic discoveries are derived almost entirely from a DLD1 colon cancer cell line overexpressing an RNA-binding-defective, acetylation-mimetic TDP-43 2KQ mutant and homozygous TDP-43 K181E iPSC-derived organoids. Both systems are far from physiological. The 2KQ mutation is a synthetic double lysine-to-glutamine mutant originally designed to mimic acetylation and disrupt RNA binding. In this study, essentially all cell-based mechanistic data on phase behavior, screens, and XPO1 effects rely on 2KQ. Yet there is no quantification of how much endogenous TDP-43 is acetylated in degenerating human neurons, nor whether a 2KQ-like acetylation state is ever achieved in vivo. It is not established that the phase behavior of 2KQ recapitulates the physiological or pathological phase behavior of wild-type TDP-43 or genuine disease-linked mutants, which may retain partial RNA binding and different post-translational modification patterns. As a result, it is difficult to know whether the modifiers identified here regulate a highly artificial 2KQ condensate or physiologically relevant TDP-43 condensates. To address this concern, the paper would benefit from quantifying endogenous TDP-43 acetylation at the relevant lysines in control and ALS/FTD patient tissue or more disease-proximal models such as heterozygous TARDBP mutant iPSC neurons, which would justify the focus on an acetyl-mimetic mutant. Key phenomena, including XPO1 dependence of phase behavior, effects of proteasome and HSP90 inhibition, and effects of splicing and translation inhibitors, should be tested for wild-type TDP-43 expressed at near-physiological levels and for one or more bona fide ALS/FTD-linked TARDBP mutants that are not acetyl mimetics. At a minimum, the authors should show that endogenous TDP-43 in neuronally differentiated cells exhibits qualitatively similar responses to XPO1 modulation, rather than exclusively relying on DLD1 2KQ overexpression.

Acetylation of endogenous TDP-43 was reported by several studies. Although it occurs at low levels under normal conditions, TDP-43 acetylation is upregulated under stress conditions (e.g. oxidative stress and proteotoxic stress) (PMID: 25556531; PMID: 28724966). Importantly, Cohen et al. reported the identification of acetylated TDP-43 in ALS patient spinal cord (PMID: 25556531), while Yu et al. showed that endogenous wildtype TDP-43 undergoes demixing when neurons were treated with either a deacetylase inhibitor or proteasome inhibitor (PMID: 33335017). These studies also show that acetylated TDP-43 is defective in RNA binding and more prone to aggregation. Furthermore, ectopic expression of acetylated TDP-43 mimetics in cells and mice induces cellular defects similar to those observed in disease models (PMID: 28724966). Thus, our findings, based on previously established TDP-43 mimetics, should provide valuable information regarding the phase regulation of a disease-relevant TDP-43 mutant. We have included more background information to justify the use of TDP-43 acetylation mimetics in the introduction.

(2) The organoid model is based on a homozygous K181E knock-in line. However, in patients, TARDBP mutations are overwhelmingly heterozygous. Homozygosity is thus a severe, arguably non-physiological sensitized background that may exaggerate nuclear RNA mis-splicing and phase defects and alter the relative contribution of cytoplasmic aggregation versus nuclear loss-of-function. In addition, it is not fully clear from this manuscript whether the structures in K181E organoids are bona fide anisosomes as defined in Yu et al. 2021, characterized by HSP70-enriched central liquid cores with TDP-43 shells and similar FRAP and fusion behavior to anisosomes in the DLD1 model. At present, the organoid section is framed as validation of "anisosome-bearing organoids," but the figures in this manuscript mainly show pTDP-43 puncta and total TDP-43 immunostaining, without detailed structural or biophysical characterization. The authors should explicitly compare heterozygous K181E/+ organoids or another heterozygous TARDBP mutant line with homozygous K181E/K181E organoids to assess whether XPO1 inhibition has similar effects in a genotype that more closely resembles patient genetics. They should provide direct evidence that the K181E condensates in organoids are anisosomes through HSP70 core immunostaining, three-dimensional reconstruction, and FRAP measurements, and clarify whether KPT-276 is acting on anisosome-like structures or more generic cytoplasmic aggregates or puncta. Without this, the leap from a DLD1 2KQ cancer cell model to human ALS/FTD-relevant neurons is not convincingly supported.

The reviewer is correct that the use of homozygous K181E organoids generates a background that is more sensitive for detecting phospho-TDP-43. The goal was to test whether XPO1 inhibition mitigates the phosphorylation of a TDP-43 disease mutant. For this purpose, we believe that our experimental setup is suitable. We agree that we should not extrapolate the result to over emphasize on its disease connection. We have revised the paper to tone down this section. We also remove the RNAseq data as it is not essential for our conclusions.

It is also noteworthy that TDP-43 disease mutations are usually loss-of-function alleles. Although heterozygous background is sufficient to induce disease phenotype in aged humans, heterozygous background in experimental settings is usually unable to generate severe defects. Thus, it is quite common to study TDP-43 disease-related defects in homozygous knockout or RNAi-mediated depletion conditions (e.g. PMID: 35197626; 41120751; 38277467).

Regarding the immunostaining signals in K181E organoids, we did not report them as anisosomes. As documented in the literature, p-TPD-43 is widely used as a marker to indicate pathological TDP-43 aggregation. P-TDP-43 is enriched in pathological aggregates in human ALS and FTD patients, colocalized with other aggregation signatures such as ubiquitin and other aggregation-prone proteins in the cytoplasm (PMID: 36008843), and is being used as a diagnostic marker for neurodegeneration (PMID: 31661037). The characterization of K181E organoid is reported in a pre-print by Zhang Q. et al., 2026 (PMID: 41292965), which is currently under revision for Science Advances. In Fig. 1I of this manuscript, we confirmed the cytosolic localization of p-TDP-43 in cells that were isolated from K181E organoids. In the current manuscript, Figure 7 is to show that nuclear export inhibition mitigates the accumulation of p-TDP-43 in a brain-like tissues. We revise the subheading and the corresponding text to avoid the confusion.

(3) The title and framing assert that "nuclear export governs TDP-43 phase transitions." However, prior studies such as Pinarbasi et al. 2018 and Duan et al. 2022 indicate that TDP-43 is not a canonical XPO1 cargo and that its export is largely passive, with active nuclear import being the dominant determinant of nuclear localization. The authors cite these studies but still position XPO1 as a central, quasi-direct regulator. The data presented are largely correlative or based on pharmacologic manipulation and overexpression in an overexpression mutant background, with no direct evidence that XPO1 engages TDP-43 in a specific, regulated manner. Even if XPO1 does not engage WT TDP-43, it could still engage the 2KQ variant, which needs to be tested.

We did not mean to conclude or imply that the regulation of TDP-43 by XPO1 is direct. In fact, we explicatively mentioned on page 8 of the original manuscript that the regulation is likely indirect and mediated by other factors. The sentence reads as “Since XPO1 does not bind TDP-43 directly (Pinarbasi et al., 2018), additional factors might link XPO1-mediated nuclear export to TDP-43 nuclear egression.”

We now add new data in Figure 6, showing that in an in vitro reconstitution assay using semi-permeabilized cells, LMB treatment significantly stabilizes anisosomes in an RNA dependent manner. This new data suggests that XPO1 inhibition leads to increased nuclear RNA availability, which indirectly favors anisosome assembly and maturation (see discussion). We believe that this new finding has provided significant new insight into how nuclear transport modulates TDP-43 phase behavior. We have revised the title, the abstract and changed the framing according to the reviewer’s suggestion.

(4) The XPO1 perturbations yield somewhat confusing phenotypes. XPO1 inhibition using Leptomycin B, KPT-276, and Verdinexor reduces anisosome number and enlarges remaining anisosomes, which remain liquid-like by FRAP recovery and fusion assays and stay nuclear. XPO1 overexpression causes fewer, enlarged puncta, but these are FRAP-impaired (gel-like) and redistribute to the cytoplasm. Thus, both decreased and increased XPO1 activity reduce anisosome number and enlarge puncta, but with opposite phase behaviors and subcellular localizations. The model presented in Figure 5L is relatively qualitative and does not resolve these issues. Moreover, XPO1 inhibition globally impairs nuclear export of many cargos and profoundly alters the nuclear environment, transcription, RNA processing, and chromatin. It is therefore difficult to conclude that the observed effects are specific to TDP-43 phase regulation as opposed to secondary consequences of broad nuclear export blockade.

The reviewer correctly summarizes our data and interpretation: XPO1 loss-of-function and gain-of-function generate opposite phenotypes regarding TDP-43 phase regulation.

Regarding the mechanism underlying XPO1-dependent TDP-43 phase regulation, as mentioned above, we developed a semi-permeabilized cell-based assay in which we used the pore-forming toxin streptolysin O to damage the plasma membrane after anisosome induction. We noticed that upon cell permeabilization and cytosol loss, anisosomes were mostly lost (Figure 6B, C). This is probably due to a reversible partition of TDP-43 into a less fluorescent soluble fraction. Supporting this idea, when permeabilized cells were incubated with cytosol plus an energy regenerating system, small puncta containing TDP-43 2KQ could be reformed in an energy dependent manner (Figure 6D, E). Interestingly, in LMB-treated cells, anisosomes remained stable despite cell permeabilization（Figure 3F). Since LMB treatment did not increase TDP-43 nuclear concentration (Supplemental Figure 1), this data suggest that nuclear export inhibition likely alter the nuclear environment to stabilize anisosomes. Indeed, when cells were permeabilized in the presence of a small RNAase, LMB-stabilized anisosomes also collapsed (Figure 6G).

We now add more discussions on the potential effect of RNA on TDP-43 phase behavior in XPO-1 inhibited cells considering these new findings.

(5) The authors show that anisosome induction depletes nucleoplasmic XPO1 signal and that mCherry-XPO1 can be seen in some TDP-43 puncta. However, antibody penetration into anisosomes is limited, so XPO1 depletion from nucleoplasm could reflect sequestration in the anisosome shell or core, but this is not demonstrated. There is no demonstration of physical interaction, even indirect interaction, between XPO1 and TDP-43 or a defined adaptor, nor identification of a specific mutant of XPO1 that selectively disrupts this putative interaction while preserving other functions. The known TDP-43 NES has been shown to be weak and not a functional XPO1-dependent NES in multiple studies. If XPO1 is acting through an adaptor that recognizes 2KQ or K181E specifically, that by itself would bring into question the generality of the mechanism for wild-type TDP-43.

We agree that our data does not demonstrate an interaction between XPO1 and TDP-43. Considering our new data (mentioned above), it is possible that the effect of anisosome induction on endogenous XPO1 localization is also mediated by RNA. We now mention more explicitly that the regulation of TDP-43 by XPO1 is likely indirect (Page 8). We have revised our paper to separate any speculative statements from the data, and also discussed the possibility of alternative interpretations.

(6) To support a mechanistic claim that nuclear export governs TDP-43 phase transitions, more targeted evidence is needed. The authors should test whether siRNA knockdown or CRISPR interference of XPO1 in the DLD1 2KQ model reproduces the effects seen with Leptomycin B and KPT-276, including FRAP and fusion phenotypes, and verify on-target effects by rescue with an siRNA-resistant XPO1 construct. They should demonstrate that canonical XPO1 cargos behave as expected under the inhibitor conditions used, as a positive control, and that the concentrations used are not grossly toxic. They should attempt to identify or at least constrain candidate adaptors that might enable XPO1-dependent export of TDP-43 through proteomic analysis of XPO1 co-purifying with 2KQ condensates or loss-of-function studies of candidate adaptors from the siRNA screen. Finally, they should test whether a TDP-43 mutant that cannot bind the proposed adaptor still responds to XPO1 manipulation.

The anisosome enlargement phenotype upon XPO1 depletion was seen in our siRNA screens, which was identified by machine-based image analyses using 6 different siRNAs. This, together with the chemical inhibition experiments, demonstrate that the phenotype is specifically caused by XPO1 inactivation.

When characterizing the effect of XPO1 inhibition on anisosome dynamics, we preferred chemical inhibitor because the effect is acute, and therefore less likely to be secondary.

Regarding the inhibitor concentration, according to the literature, Leptomycin B was commonly used at 50-200 nM. We chose 200 nM to ensure a quick and complete inhibition of XPO1-mediated nuclear export (see Figure 3 in PMID: 9628873). This dose is also well tolerated by our cells.

We did not suggest any specific adaptor that mediates XPO1 interaction with TDP-43. Whether there is an adaptor, and if so, the identity of such adaptor is out of the scope of this study. We revise our paper on page 8-9 to clarify these points.

(7) Even with these data, what is currently shown is that global modulation of nuclear export capacity can alter the phase behavior and localization of a highly overexpressed RNA-binding-defective TDP-43 mutant and of K181E in organoids. This is important, but it is weaker than asserting that XPO1 directly governs TDP-43 phase transitions in physiological contexts. The title, abstract, and Discussion should be tempered to reflect that nuclear export is one of several pathways, alongside RNA splicing, translation, and proteostasis, that influence TDP-43 phase states in this model, and that the specific mechanism and cargo relationship between XPO1 and TDP-43 remain unresolved and may be indirect.

We have revised the title, abstract, and main text to temper our conclusions.

(8) The authors conclude that cytoplasmic TDP-43 aggregation plays only a modest role in TDP-43 proteinopathies because in homozygous K181E organoids, chronic KPT-276 treatment almost abolishes cytoplasmic pTDP-43 puncta, yet bulk RNA-seq shows only a relatively small fraction of dysregulated genes are rescued. There are several issues with this inference. Relying primarily on pTDP-43 antibody staining to define cytoplasmic TDP-43 aggregation is limiting. pTDP-43 antibodies label only phosphorylated species and may miss non-phosphorylated, oligomeric, or amorphous TDP-43 species that could still be toxic. Different pTDP-43 antibodies vary in epitope accessibility depending on aggregate conformation and subcellular location. More sensitive approaches, such as high-affinity TDP-43 RNA aptamer probes developed by Gregory and colleagues, biochemical fractionation for SDS-insoluble and urea-soluble TDP-43, and filter-trap assays, would provide a more quantitative assessment of cytoplasmic aggregation and its reduction by KPT-276. Without these, it is not safe to assume that cytoplasmic aggregation has been eliminated, as opposed to one antigenic subclass.

We agree with the reviewer that p-TDP-43 may not represent all aggregate species. However, p-TDP-43 antibodies detect the pathologically validated species tightly associated with TDP-43 proteinopatheis. In human ALS and FTD-TDP tissues, cytoplasmic inclusions are strongly immunoreactive for phosphorylated TDP-43 (typically S409/410, as detected here). Additionally, p-TDP-43 immunohistochemistry is a routine diagnostic criterion in neuropathology. For these reasons, we believe that the observation that inhibition of XPO1 significantly reduces p-TDP-43 is a significant finding, as it suggests that inhibition of nuclear transport may rescue TDP-43 proteinopathy. We revised the text on page 9 to better explain the significance of p-TDP-43 staining.

(9) The treatment window, spanning from day 87 to 122 with 20 nanomolar KPT-276, may be too late or too mild to reverse entrenched nuclear RNA-processing defects, even if cytoplasmic inclusions are cleared. Once widespread cryptic exon inclusion and alternative polyadenylation misregulation are established, many downstream changes may become self-sustaining or only partially reversible. Moreover, XPO1 inhibition will massively rewire nucleocytoplasmic transport of many transcription factors, splicing factors, and RNA-binding proteins. Thus, the lack of full transcriptomic rescue cannot be cleanly interpreted as evidence that cytoplasmic aggregates are only modest contributors. It may instead reflect that nuclear dysfunction is primary and XPO1 inhibition does not correct, and may even exacerbate, certain nuclear defects.

We agree with the reviewer that the lack of rescue may be caused by some technical issues. We have removed the RNAseq data and the related texts since it is not essential.

(10) To support a causal statement about the modest contribution of cytoplasmic aggregates, one would want more direct measures of neuronal health and function, such as cell death, neurite complexity, synaptic markers, and electrophysiology before and after KPT-276, not only transcriptomics. A way to selectively reduce cytoplasmic aggregation without globally inhibiting nuclear export would allow comparison of outcomes.

We have removed the discussion regarding the role of cytoplasmic aggregates in disease.

(11) Given these caveats, the concluding statements that cytoplasmic TDP-43 aggregation is only a modest contributor should be substantially softened. A more defensible interpretation is that in this homozygous K181E organoid model, chronic global XPO1 inhibition reduces pTDP-43-positive cytoplasmic puncta but only partially normalizes the steady-state transcriptome, suggesting that persistent nuclear RNA-processing defects and other pathways continue to drive pathology.

We agree with the review and have removed the RNAseq part.

(12) The screens are a major strength but need more rigorous validation for key hits, especially nuclear transport factors. For the siRNA screen, hits are filtered by anisosome number per nucleus, but there is no direct demonstration in the main text that XPO1 or CSE1L knockdown is efficient at the messenger RNA or protein level. For the highlighted genes, Western blot or quantitative polymerase chain reaction validation and phenotypic rescue would strengthen confidence. For small-molecule hits, it is not systematically shown that anisosome modulation is independent of changes in total TDP-43 2KQ expression or gross toxicity. Translation inhibitors are tested for this, but for many other hits, including proteasome, HSP90, and kinase inhibitors, expression and general nuclear structure should be monitored. Given the reliance on anisosome count as a readout, secondary screens that specifically distinguish changes in TDP-43 expression levels, changes in nuclear morphology or cell cycle, and specific changes in anisosome phase behavior, including FRAP and fusion for top hits, would greatly increase interpretability.

For the siRNA screen, each positive hit was confirmed by two rounds of screen with 6 independent siRNAs in total. Although we did not validate the knockdown efficiency due to the large number of hits, we routinely include a positive siRNA control in our study (Cell death siRNA), which targets several essential gene. Transfection efficiency was controlled by measuring cell viability after knocking down of these genes. In addition, the identification of XPO1 as a positive regulator of TDP-43 phase behavior was independently validated by our chemical genetic screens with three XPO-1 inhibitors. We feel confident that XPO1 is a key modulator of TDP-43 phase behavior.

For chemical treatment experiments, the anisosome fusion phenotypes could be detected as early as 5 h post treatment. Given the relatively short treatment, we do not expect a significant change in protein level or toxicity. To alleviate this reviewer’s concern, we performed an immunoblotting experiment to measure the total TDP-43 protein levels in drug-treated cells. Except for VLX, we did not detect any significant changes in the level of TDP-43 after drug treatment (Supplemental Figure 1).

(13) The classification of condensates as liquid versus gel-like or solid is based almost entirely on FRAP recovery or lack thereof. While FRAP is appropriate, interpretations could be made more robust by including half-region-of-interest bleach controls and assessing mobile fractions and recovery kinetics more quantitatively across conditions. Complementing FRAP with other phase-behavior assays such as sensitivity to 1,6-hexanediol, shape relaxation after deformation, and coarsening behavior over longer timescales would strengthen the analysis. At present, some assignments, such as that XPO1 overexpression drives a gel-like transition, are reasonable but somewhat qualitative.

In this study, we used two types of FRAP assays. We either bleached TDP-43 within anisosomes or bleached the surrounding TDP-43 molecules（Figure 2）. The two complementary methods yield consistent results that allow unambiguously distinguish between TDP-43 LLPS state and gel-like condensation.

In XPO1-related experiments, the two types of condensates formed by TDP-43 2KQ can be distinguished by several features including their subcellular localization, shape, and the fluorescence recovery kinetics. We feel that these combined data clearly segregate these puncta into two distinct types of assemblies. The proposed half-region-of-interest bleach is technically challenging for small anisosomes under normal conditions. However, whenever possible, (e.g. anisosomes enlarged by Leptomycin B), we did perform both whole anisosome bleach and partial bleach (Figure 5D, I). Both assays demonstrate that TDP-43 in these enlarged anisosomes is highly mobile.

(14) For the Leptomycin B and KPT-276 experiments in cells and organoids, it would be important to confirm that canonical XPO1 cargo proteins accumulate in the nucleus and that the concentrations used are within a range that is not overtly toxic over the experimental timeframe. Assessing nuclear morphology, chromatin condensation, and general transcriptional activity through global RNA synthesis or key reporter genes would ensure that observed effects are not secondary to severe global nuclear export collapse.

In Leptomycin B treatment experiments, we carefully chose a dose that was previously validated (see Figure 3 in PMID: 9628873). Based on our DAPI staining, the nuclear morphology appears normal with no abnormal chromosome condensation (Figure 5A). Additionally, in cell line-based experiments, the effect of Leptomycin B on anisosomes was detected 6-8 hours post treatment. The change in global protein synthesis because of RNA changes should be relatively minor at this stage. Indeed, our new immunoblotting experiment showed that LMB treatment did not affect TDP-43 protein level (Supplemental Figure 1). Most importantly, the in vitro semi-permeabilized assay demonstrates a direct role for RNA in stabilizing anisosomes.

(15) In the organoid section, it is not clear how many independent iPSC clones and organoid batches were used per condition, nor whether batch effects were assessed in the bulk RNA-seq analysis. This should be fully specified and ideally controlled with isogenic wild-type and K181E clones. For transcriptional rescue, it is important to know whether the changes in wild-type organoids treated with KPT-276 are negligible. A direct wild-type comparison with or without KPT-276 is important to disentangle general drug effects from K181E-specific rescue. More detailed quantification of total TDP-43 and pTDP-43 in both nuclear and cytoplasmic fractions, including biochemical fractionation if possible, would strengthen the assertion that KPT-276 specifically reduces cytosolic pTDP-43 aggregates while sparing nuclear TDP-43.

The organoid experiment was performed with two batches per condition to reduce the effect of batch variation. The wildtype cells and K181E mutant are derived from the same genetic background. This information is now included in the method section on page 14. Given the criticisms by review 1 and 2 on the RNAseq data, we have removed this non-essential data.

(16) Beyond the core issues above, several additions could greatly enhance the impact. The manuscript currently emphasizes XPO1, but the genetic and chemical data clearly implicate RNA splicing, translation, and proteostasis as equally strong or stronger regulators of TDP-43 phase states. A more integrated model that explains how these pathways intersect, for example, how splicing factor availability, ribosome loading, and proteasome capacity co-govern anisosome nucleation, growth, and hardening, would be valuable.

We now discuss a new model in discussion based on our new Figure 6, which integrates the role of RNA splicing and nuclear transport in TDP-43 phase regulation on page 10. We agree with the reviewer that other questions are also important for future studies.

(17) A key unresolved question is whether XPO1 is acting directly on TDP-43, or instead primarily regulates anisosomes by exporting other factors that more proximally control TDP-43 phase behavior. Given that TDP-43 is not a canonical XPO1 cargo and prior work indicates that its nuclear export is largely passive, it seems at least as plausible that XPO1 inhibition alters the nuclear concentration or localization of splicing factors, RNA-binding proteins, chaperones, or other modifiers identified in the screens, and that changes in these proteins secondarily reshape anisosome dynamics. In other words, XPO1 may be exporting a more direct regulator of anisome formation and hardening, rather than exporting TDP-43 itself in a specific, regulated way. The current data do not distinguish between these possibilities. Systematic identification of XPO1-dependent cargos that colocalize with or biochemically associate with anisosomes, combined with targeted perturbation of their nuclear export, would be needed to determine whether the relevant XPO1 substrate in this system is actually TDP-43 or an upstream modulator of its phase behavior.

As discussed above, our new data regarding the role of RNA in TDP-43 phase regulation should alleviate this concern, although we cannot exclude the possible involvement of splicing factors in this process. We also clearly state that there is no evidence to support a direct interaction between TDP-43 and XPO1 on page 8.

(18) Testing whether identified modifiers converge on nuclear TDP-43 concentration would be informative. Since phase separation is concentration-dependent, measuring nuclear versus cytoplasmic TDP-43 levels across key perturbations, including splicing inhibition, translation inhibition, proteasome inhibition, HSP90 inhibition, and XPO1 modulation, would help determine whether modifiers mainly work by changing nuclear TDP-43 concentration or by altering interaction networks and the material properties of condensates.

In the newly performed immunoblotting experiment, we measured the TDP-43 levels in drug-treated cells but found no effect by most drugs (Supplemental Figure 1).

(19) Examining other ALS-relevant RNA-binding proteins would be valuable. Given the role of XPO1 and other hits, it would be informative to briefly test whether similar principles apply to FUS, hnRNPA1, or other ALS-relevant RNA-binding proteins in the same cellular context, to argue for generality versus TDP-43-specific idiosyncrasies of the 2KQ system.

We agree that this is an important issue but we feel the proposed experiments are beyond the scope of the study.

(20) The Introduction sometimes implies that anisosomes are common and well-established intermediates en route to pathology. It would be helpful to more clearly state that, to date, anisosomes are primarily observed in overexpression and mutant systems and have not yet been unequivocally demonstrated in human patient tissue. The link between PDGFRβ, PAK4, GSK-3β, and YAP and TDP-43 phase dynamics is intriguing but only briefly mentioned. The authors should either expand on this or tone down the emphasis in the Results section.

We have revised the introduction and added the following sentence on page 4. “The 2KQ-containing anisosomes, observed mostly in the nucleus under overexpression conditions, have not been validated in human patient samples.”

(21) In the organoid methods, the authors should consider clarifying whether doxycycline is continuously used, which might alter TDP-43 expression and nuclear transport in a non-negligible way.

The organoid model does not involve protein overexpression or doxycycline treatment. We measured endogenous p-TDP-43, which is why we feel this experiment is very significant. Unlike many other p-TDP-43 detection studies that rely on TDP-43 overexpression or exposing cells to excess stressors, we could detect substantial p-TDP-43 in 3D organoids grown under normal conditions, whereas the same cells grown and differentiated in 2D culture do not show p-TDP-43 (Zhang Q. et al., BioRxiv 2025).

(22) For statistical methods, it would be beneficial to indicate whether multiple-comparison corrections were applied for the many FRAP, anisosome count, and size comparisons beyond DESeq2 internal corrections for RNA-seq.

We have added more statistical information to the figure legends.

(23) Some figure legends could more clearly indicate whether the images shown are single z-planes or maximum intensity projections and how the thresholding for anisosome detection was performed.

We revised the figure legends to include this information. As for anisosome detection, because they are so obvious, standard thresholding combined with automated counting was sufficient to identify them.

(24) In its current form, the manuscript contains an impressive set of screens and some nicely executed imaging of TDP-43 condensates, highlighting nuclear export among other pathways as a modulator of TDP-43 phase behavior. However, the physiological relevance is undercut by heavy reliance on an acetylation-mimetic, RNA-binding-defective TDP-43 mutant and a homozygous K181E organoid model. The mechanistic link between XPO1 and TDP-43 remains largely inferential and partly at odds with prior work. The conclusion that cytoplasmic TDP-43 aggregation is only a modest contributor to disease is not firmly supported by the available data.

We agree with the reviewer that the strength of the study is our unbiased approach that identifies pathways capable of modulating TDP-43 phase behavior. In the revised paper, we included several experiments using an in vitro semi-permeabilized cell system to further dissect the role of nuclear export in TDP-43 phase separation. We believe that these new results should provide significant mechanistic insight that links nuclear export and RNA transcription and splicing to TDP-43 phase regulation. Additionally, we have revised our paper carefully to discuss the physiological relevance and the limitation of our study.

(25) With substantial additional mechanistic work, particularly around XPO1, rigorous validation in more physiological TDP-43 contexts, more sensitive detection of cytoplasmic TDP-43 aggregates, and a tempering of the central claims, this study could make a meaningful contribution to understanding how nucleocytoplasmic transport and other cellular pathways influence TDP-43 phase transitions and aggregation. The work should be reframed as an important screening study that identifies nuclear export as one among several cellular processes that modulate TDP-43 phase behavior in a model system, rather than as a definitive demonstration that nuclear export governs pathological TDP-43 aggregation in disease.

We now reframe the study as an important screening study that identifies nuclear export among several other pathways as modulators of TDP-43 phase behavior. We also propose a model that links RNA splicing to nuclear export in TDP-43 phase regulation.

Reviewer #2 (Public review):

Summary:

This manuscript addresses an important and timely question in TDP-43 biology by systematically identifying regulators of TDP-43 anisosome formation, with a particular focus on nuclear export via XPO1. Using a combination of unbiased chemical screening, genetic perturbation, and advanced imaging approaches, the authors propose that inhibition of nuclear export modulates the abundance and biophysical properties of TDP-43 anisosomes. The study is conceptually innovative and has potential relevance for neurodegenerative diseases characterized by TDP-43 pathology. However, significant concerns regarding experimental controls, reporting transparency, and model translatability currently limit the strength of the conclusions and the interpretability of several key findings.

We thank the reviewer for acknowledging the significance and innovation of our study.

Strengths:

(1) The study employs an unbiased, hypothesis-free compound screen to identify regulators of TDP-43 anisosome formation, which is a major strength and reduces confirmation bias.

(2) The authors combine chemical and genetic screening approaches, providing orthogonal validation of key pathways and increasing confidence in the biological relevance of top hits.

(3) The focus on biophysical properties of TDP-43 assemblies, assessed through imaging and FRAP, moves beyond simple presence/absence of aggregates and provides mechanistic insight into the biophysical states of TDP-43.

(4) The use of multiple experimental modalities, including live-cell imaging, FRAP, pharmacological perturbation, and transcriptomic analysis, reflects a technically sophisticated and ambitious study design.

(5) The authors attempt to extend findings beyond immortalized cancer cell lines by incorporating organoid models, demonstrating awareness of disease relevance and translational importance.

Overall, the manuscript is clearly written and logically structured, making complex experimental workflows accessible and the central hypotheses easy to follow.

Weaknesses:

Despite its strengths, the manuscript has several major limitations that affect data interpretation and confidence in the conclusions.

(1) Lack of appropriate controls for overexpression experiments:

A central concern is the absence of proper controls for TDP-43 and XPO1 overexpression. Prior studies (including those cited by the authors, Archbold et al.2018) show that overexpression of WT TDP-43 alone is toxic to neurons. Thus, the experimental system itself may induce anisosome formation independently of the mechanisms under study. Similarly, XPO1 overexpression lacks a suitable control (e.g., mCherry alone or mCherry fused to a protein known to be independent of TDP-43). The near-complete colocalization of XPO1 with TDP-43 anisosomes upon overexpression raises the possibility that these structures reflect non-physiological protein accumulation rather than regulated assemblies.

As mentioned in our response to reviewer 1, point 1, we have added more discussions to justify the use of acetylation mimetics in our study. We agree with the reviewer that these large puncta (both anisosomes and gel-like structures) likely resulted from TDP-43 overexpression. Nevertheless, in a titration experiment done by Yu et al. 2020 (PMID: 33335017), they showed that ectopic TDP-43 undergo demixing even at concentrations lower than endogenous TDP-43, although the demixed puncta were very small. Their result suggested that overexpression per se does not change TDP-43 phase behavior, only enlarge the demixed TDP-43 structures, which is necessary for our screen and imaging-based characterization.

For XPO1 overexpression, we have done the mCherry alone control but due to space limit in Figure 5, we did not include it. We now include the data in Supplemental Figure 4. This figure shows that overexpression of mCherry did not change TDP-43 localization or anisosome structures.

(2) Insufficient experimental and analytical transparency:

The manuscript frequently lacks clear reporting of experimental details. In multiple figures, the stated number of independent experiments does not match the number of data points shown, making it difficult to assess statistical validity. Concentrations used in the compound screen are not clearly defined, nor is it stated whether multiple concentrations were tested. It is unclear how many wells, cells, or independent cultures were analyzed. The criteria used to reduce 1,533 screening hits to 211 candidates via STRING analysis are not explained. Knockdown and overexpression efficiencies are not reported.

We apologize for these omissions. We have added more experimental details to the figure legends and the method. For the imaging experiments, data points reflect randomly selected individual cells imaged in 2-3 independent biological repeats. This is now stated in the figure legends. For chemical screens, we screened against NCATS libraries was first done at top concentration (10 mM) to ensure inhibitory efficacy for all potential hits. In the follow-up validation study, we validated the top hits using a series of concentrations, as shown in Figure 1B. Drug concentrations are provided in Figure 2A, 4A, C, E, F, 5A-D, F, Figure 6F, G, Figure 7A)

We explain the STRING analysis in more detail now. Basically, STRING is a protein-protein interaction network that reports all potential interactions between any proteins in human proteome. Given the potential off-target effect of siRNA, we assume that if the screen identifies multiple components of a protein interaction network or pathway, the result is more likely to be real.

We did not check XPO1 knockdown efficiency in high through-put screens (HTS) for several reasons. Firstly, the large number of positive hits makes it impossible to check knockdown efficiency for all of them. Secondly, the effect of XPO1 knockdown on anisosomes was seen with 6 different siRNAs in two rounds of screens. Thirdly, in the HTS protocol, we routinely included a transfection control (siRNAdeath) to control transfection efficiency. We would only process the data if siRNAdeath control killed > 90% of the cells. Lastly, the XPO1 knockdown result was independently validated by small molecule inhibitors. For TDP-43 overexpression, the study by Yu and colleagues suggested that the expression is more than 20-fold higher than endogenous TDP-43, but they showed that anisosome formation is not an artifact of protein overexpression. When the expression level was titrated down, they could still detect anisosomes.

(3) RNA-seq concerns:

The RNA-seq experiments are particularly problematic. The number of biological replicates per condition is not stated, and heatmaps suggest that only one sample per group may have been used, which would preclude statistical analysis. No baseline comparison between WT and mutant TDP-43 is shown. Given that TDP-43 is an RNA-binding protein, splicing analyses would be far more informative than gene expression alone, yet no splicing data are presented. Moreover, nuclear retention of TDP-43 does not preclude nuclear aggregation, which may still impair its splicing function.

We apologize for the lack of clarity regarding the RNA-seq design. For each condition, organoids of two independently differentiated batches were treated in triplicate. What we showed before was averaged expression levels. We pooled the organoids of the same treatment from the two batches to reduce the impact of batch variation.

Given the criticisms from both reviewers 1 and 2 on the limited interpretation power of the RNAseq study, we have removed this data from the revised manuscript.

(4) Limited translatability to neuronal biology:

All anisosome analyses are performed in a cancer cell line, raising concerns about relevance to post-mitotic neurons. While organoids are used as a secondary model, the assays performed do not overlap with those used in cancer cells, making it difficult to assess whether anisosome-related mechanisms are conserved. Neuronal toxicity, a critical outcome given known TDP-43 biology, is not assessed. Prior work has shown that WT TDP-43 overexpression alone is toxic to neurons, yet this is not addressed.

We agree with the reviewer that the model used in this study is not directly relevant to neurodegeneration. However, as pointed out by the reviewer, neurons are much more sensitive to TDP-43-associated toxicity. By contrast, the cell line used in this study can tolerate TDP-43 overexpression with no detectable cytotoxicity. This feature makes it feasible to evaluate how different cellular processes modulate TDP-43 phase behavior without the confounding effect from cytotoxicity. Notably, the processes identified by our screens are all house-keeping pathways that are conserved in neurons. Thus, we believe that the reported findings are likely applicable to neurons. That being said, we have revised our paper to ensure that we don’t overstate the clinical relevance of our work.

(5) Conceptual and interpretational gaps:

The authors quantify anisosome number but also report conditions in which anisosome number decreases while size increases. The biological interpretation of larger anisosomes is not discussed, and whether this reflects improvement or worsening of pathology is unclear. Compounds targeting the same mechanism (e.g., nuclear export inhibition) are inconsistently used across experiments (KPT compounds, verdinexor, leptomycin B), raising concerns about reproducibility. In organoids, the experimental paradigm shifts to long-term treatment (35 days vs. 16 hours), further complicating interpretation.

We thank the reviewer for these critical points. As pointed out by the reviewer 1 in point 4 above, we do not have evidence to establish a convincing correlation between the size of anisosomes and clinical phenotypes. Regarding the use of different drugs for different experiments, the initial screen identified KPT and Verdinexor because they are investigational drugs, but Leptomycin B was not in our library. In the follow-up studies, we switched to Leptomycin B because 1) it is highly potent and specific; 2) it was better characterized and more commonly used as inhibitors of XPO1 according to the literature. However, for the organoid study, we had to switch back to KPT because of the toxicity issue associated with long-term application of Leptomycin B.

(6) Overinterpretation of rescue effects:

Although the authors state that they aim to test whether nuclear export inhibition rescues neuronal defects, no functional neuronal readouts are provided (e.g., viability, morphology, axon outgrowth, or electrophysiological measures). RNA-seq alone is insufficient to support claims of rescue.

Our interpretation of the RNA-seq data was that the rescue effect by nuclear export inhibition was limited and probably insignificant. Given that this negative data is not conclusive, we have removed it from the revised manuscript.

(7) Finally, the model does not appear to exhibit cytosolic TDP-43 aggregation at baseline. It remains unclear whether longer induction would produce cytosolic gel-like assemblies and whether these would be prevented by nuclear export inhibition. Long-term data are shown only in organoids, yet anisosome formation is not assessed there.

The expression system used in the study reaches a steady state after 24 h of induction. Prolonged expression up to 48 h did not alter the number of anisosome, nor does it change TDP-43 phase behavior. We now clarify this point on page 4.

Reviewer #3 (Public review):

Summary:

TDP-43 proteinopathy is broadly found in neurodegenerative diseases. This manuscript investigates how nuclear export influences the biophysical properties of TDP-43. The authors use a combination of chemical screening and genome-wide siRNA screening to identify pathways that modulate TDP-43 liquid-to-solid transitions. Overall, the study employs a broad array of approaches and addresses an important question in TDP-43 pathobiology. The identification of nuclear export as a central regulator is compelling and conceptually aligns with the emerging view that TDP-43 nucleocytoplasmic trafficking is a major defect in neurodegeneration.

Strengths:

This work integrates chemical and genetic screening to identify novel modifiers. The candidates were validated in both reporter cell lines and iPS-differentiated organoids. The findings support the nucleocytoplasmic transport is important for the biophysical properties of TDP-43.

We thank the reviewer for acknowledging the significance and strength of our study.

Weaknesses:

The mechanisms underlying the connection between nuclear export and phase transition need further clarification. Broader consequences of XPO1 inhibition are not addressed.

We agree that our previous manuscript did not address how nuclear export inhibition affect TDP-43 phase behavior. As discussed in our paper, we proposed that the effect of nuclear export inhibition on TDP-43 phase separation is likely indirect. The most likely scenario is that inhibition of nuclear export changes the nuclear environment over time, which affects TDP-43 phase separation. We have tried to isolate nuclear extracts from control and LMB-treated cells and used mass spectrometry to identify proteins that are differentially present in the nucleus. However, knockdown of the identified top candidates did not abolish LMB-induced phase alteration (not shown). Considering our observation that RNA splicing is another modulator of TDP-43 phase behavior, we reasoned that it is possible that it is the combined change of RNA and protein composition in the nucleus that alters TDP-43 phase behavior. In new experiments presented in Figure 6, we now used a semi-permeabilized in vitro system to demonstrate that LMB treatment stabilized anisosomes in an RNA-dependent manner (see response to point 4 by reviewer 1). This new data allows us to propose a new model that link RNA splicing and nuclear export in TDP-43 phase regulation (Discussion).

Recommendations for the authors:

Reviewer #2 (Recommendations for the authors):

(1) Include appropriate controls for all overexpression experiments. In particular, overexpression of WT TDP-43 alone and suitable tag-only controls (e.g., mCherry alone or mCherry fused to a protein unrelated to TDP-43/XPO1) should be included to control for aggregation driven by non-physiological protein levels.

In Supplemental Figure S4, we included a tag-only control, which shows that mCherry alone does not affect the localization of XPO1, neither did we see mCherry co-localizes with TDP-43.

Since WT TDP-43 itself does not form anisosome and because the goal of the study was to test how anisosome dynamics is affected by various conditions, we did not repeat our experiments with WT TDP-43.

(2) Address whether TDP-43 anisosomes form under endogenous or near-physiological expression levels. If possible, include experiments using lower expression systems or endogenous tagging to demonstrate that anisosome formation is not solely an overexpression artifact.

As mentioned above, in a titration experiment done by Yu et al. 2020 (PMID: 33335017), they showed that ectopic TDP-43 undergoes demixing even at concentrations lower than endogenous TDP-43, although the demixed puncta are small. Their result suggested that overexpression per se does not change TDP-43 phase behavior. Instead, it only enlarges the demixed TDP-43 structures, which is necessary for our screen and imaging-based characterization.

(3) Clearly define biological versus technical replicates throughout the manuscript and report exact n-numbers for all experiments in figure legends and/or methods. Resolve discrepancies between stated and displayed n-numbers (e.g., figures showing more data points than the number of independent experiments reported). Further, include how data points were defined (e.g., cells, fields of view, wells).

We now state clearly the biological repeats in figure legends. We did not use N number to specify technical replicate. The discrepancy between the stated N number (biological repeats) and the data points is because for imaging experiments, data points usually represent single cells collected from 2-3 biological replicates (N=2 or 3). Data points are now clearly defined in the figure legends (anisosome, cell, imaging field, or independent experiment).

(4) The authors state that they identified a list of compounds that reduced anisosomes. Please clarify how the threshold was determined: Was this a statistical analysis or a specific threshold that has been used?

For both siRNA screen and chemical genetic screen, we calculated the Z-score and used Z-score>2 as a cutoff. This is mentioned in the method.

(5) Provide a complete list of compounds used in the chemical screen, including concentrations tested and whether multiple doses were evaluated.

As mentioned above, the initial screen was done with just one concentration (10 mM). Identified positive hits were re-tested with multiple doses as shown in Figure 1. The compounds are from a commercial library (LOPAC R1280, Sigma #LO4200). The list of compounds can be found at vender’s website.

(6) Clearly explain the criteria used to reduce the initial 1,533 screening hits to 211 candidates following STRING analysis, including cutoffs and prioritization logic.

We now explain that the Z-score was used to further narrow down the hit (page 6). Additionally, we provide an explanation on how we use STRING to further narrow down the list. The sentence reads as “To further narrow down the list, we performed a STRING protein network analysis based on the assumption that a protein interaction network bearing multiple positive hits would be more likely to be a true effector.”

(7) Report knockdown and overexpression efficiencies for all genetic perturbations used in the study.

For TDP-43 overexpression, the study by Yu and colleagues suggested that the stable cell line expresses 20-fold more TDP-43 than endogenous one, but they showed that anisosome formation is not an artifact of protein overexpression. When the expression level was titrated down, they could still detect anisosomes (Yu, H. et al., Science 2021). For knockdown efficiency, since the screen used 6 different siRNAs for each identified target (a few hundred), it is technically challenging to validate the knockdown efficiency of each siRNA by conventional qRT-PCR. To control knockdown efficiency, we transfected cells in parallel with siRNA-death that contains a mixture of siRNAs targeting several essential genes (Qiangen, #1027299). We would only process the data if siRNAdeath control killed > 90% of the cells, indicating good knockdown efficiency.

(8) Clarify the biological interpretation of changes in anisosome size versus number, particularly in conditions where fewer but larger anisosomes are observed. Discuss whether larger assemblies are hypothesized to be protective, neutral, or deleterious.

Live cell imaging was used to dissect why cells treated with certain drugs such as XPO1 inhibitors have fewer but larger anisosome. Figure 5F shows that this is caused by the fusion of small anisosomes. Our data does not suggest that the size of anisosomes can differentiate between protective or deleterious state, but rather it is the LLPS state and subcellular localization of these assemblies that may play a more critical role in determining whether TDP-43 forms deleterious protein aggregates. The discussion is on page 10.

(9) Specify whether all anisosomes induced by XPO1 overexpression were gel-like or whether this applied only to a subset. If only a subset was affected, please provide quantifications, otherwise state clearly that all anisosomes in XPO1 overexpression were gel-like.

All TDP-43 puncta mislocalized to the cytoplasm in XPO1-overexpressing cells are gel-like because the FRAP experiment in Figure 5I was done with randomly selected TDP-43 puncta mislocalized to the cytoplasm.

(10) Clarify which anisosomes (nuclear vs cytosolic; gel-like vs non-gel-like) were selected for FRAP analyses in Figure 5I.

For Figure 5I, the control anisosomes in untreated cells are nuclear while under mCh-XPO1 expressing condition, only those in the cytoplasm were randomly selected for photobleaching.

(11) The translatability of the conclusion based on cancer cell lines to brain organoids is not convincingly shown and could be strengthened by including additional assessment of anisosomes. While this might not be feasible in 3D cultures, the authors could alternatively use 2D cultured neurons to perform the same assays as performed in the cancer cell line. Additionally, the same treatment strategy should be applied. The reasoning for increasing treatment to 35 days in the organoids is unclear.

In another manuscript that is currently under revision, we compared 2D iNeuron culture with 3D organoids. A pre-print is available at https://www.biorxiv.org/content/10.1101/2025.11.09.687455v1.full. In this study, we found that endogenous TDP-43 K181E mutant do not undergo phosphorylation-dependent transition to aggregate in 2D cultures. Only when these cells were grown into 3-D organoids, TDP-43 phosphorylation could be detected. (see supplemental Fig. S1c, d in https://www.biorxiv.org/content/10.1101/2025.11.09.687455v1.full). Thus, it is not possible to repeat the experiments in this study in 2D iNeuron cultures. We agree with the review that there is a gap between the study using the cancer cell line and the use of K181E iPSC-derived 3D organoids. We have toned down our conclusions throughout the text.

(12) Address neuronal vulnerability explicitly by assessing toxicity, viability, or functional neuronal readouts, particularly given prior reports that WT TDP-43 overexpression alone is neurotoxic.

We agree that this is an important point, but the main goal of this study was to dissect the cellular pathways/mechanisms that govern TDP-43 phase separation. We feel that the requested experiments are beyond the scope of the current study.

(13) Clearly state the number of biological replicates used for each RNA-seq condition. Establish baseline transcriptional differences between WT and mutant TDP-43 prior to assessing the effects of nuclear export inhibition. Include PCA plots and heatmaps, including all samples.

As mentioned above, we have decided to remove the RNAseq data from the manuscript to save room for new results.

(14) Given the role of TDP-43 as an RNA-binding protein, consider including splicing analyses to assess whether nuclear export inhibition preserves or disrupts TDP-43-dependent RNA processing.

We thank the reviewer for this suggestion. However, we feel that the proposed experiments are beyond the scope of the current study.

(15) Improve clarity of transcriptomic visualizations (e.g., GO-term plots) and explicitly define all group labels used (e.g., Group A vs Group B).

We have removed the RNAseq data.

(16) Ensure consistent use of disease terminology (ALS vs FTD) throughout the manuscript, e.g., lines 222 and 244.

We have checked the usage of these terms to make sure they are accurately used.

(17) Correct figure and axis labeling errors (e.g., Figure 3A x-axis range).

Figure 3A indicates the Z score distribution of the entire human genome. As stated on page 6, 21,404 genes were targeted.

(18) Avoid overstatements in the Discussion that are not directly supported by the presented data, particularly regarding the interpretation of proteasome inhibition and gel-like anisosome states.

We have revised our discussion substantially to tone down our conclusions.

(19) Clarify the rationale for switching between different nuclear export inhibitors across experiments and discuss whether results were consistent across compounds.

In the acute experiments down with the cancer cell line, we used LMB because it is potent and well characterized. In organoid experiment, we switched to KPT-276 because it is better tolerated by organoids, especially during longer treatment.

Reviewer #3 (Recommendations for the authors):

Major concerns that require clarification or further strengthening:

(1) The connection between nuclear export and liquid-solid phase transition is not clear. The 2KQ mutant forms nuclear anisosomes. The manuscript does not provide data about its nuclear-cytoplasmic distribution normally, nor how the distribution is changed upon nuclear export inhibition or enhancement. In Figure 5I, it is unclear whether the anisosomes are in the nucleus or cytoplasm. The dynamics of nuclear vs cytoplasmic anisosomes should be measured separately. What is the mechanism that promotes nuclear export and changes the dynamics, especially nuclear anisosomes?

As mentioned by the reviewer, the 2KQ mutant forms anisosomes only in the nucleus. This was documented in Yu, H. et al., Science 371 (2021), and also shown in our Figure 4A, F, Figure 5A. Figure 5A also shows that nuclear export inhibition does not change anisosome localization, only making them bigger while reducing the numbers. For Figure 5I, the control anisosomes in untreated cells are nuclear while under mCh-XPO1 expressing condition, only those present in the cytoplasm were randomly selected for bleaching.

(2) Figure 5J, no obvious XPO1 is sequestered to anisosomes, as described in lines 208-209.

Unlike Figure 5G, this experiment studied the localization of endogenous XPO-1 by immunostaining. As discussed in Yu et al., Science 371 (2021), proteins inside anisosomes could not be stained by antibodies due to an accessibility problem. This explains why we could only detect reduced XPO1 after anisosome induction.

(3) Figure 6A, the localization of phosphor-TDP-43 is not clear. And it is not clear what cell types contain the aggregates. Higher-resolution images need to be included. The mechanism by which XPO1 inhibition reduces TDP-43 aggregation requires further validation. It remains unclear whether it is directly mediated through altered nucleocytoplasmic transport of TDP-43.

We agree that it is technically challenging to visualize the precise subcellular localization of p-TDP-43 in 3D organoids. In the manuscript that reports the characterization of the 3D organoids, we dissociated cells from the 3D organoids by trypsin digestion and plated them out in 2D before immunostaining and imaging. We could clearly see p-TDP-43 co-localizes with the neuronal marker TUJ1 and is localized outside of nucleus (see figure 1 of https://www.biorxiv.org/content/10.1101/2025.11.09.687455v1.full)

In the newly added Figure 6, we used a semi-permeabilized cell system to dissect the phase separation dynamics of TDP-43 2KQ in cells treated with the nuclear export inhibitor LMB. Our data suggests that nuclear export inhibition alters the nuclear environment, making it more favorable for the liquid phase of TDP-43. This is dependent on nuclear RNA.

(4) XPO1 controls the export of numerous essential proteins, and its inhibition can produce broad, potentially toxic effects unrelated to TDP-43. The manuscript should include a discussion of these off-target consequences.

We thank the reviewer for this point. Given the new data in Figure 6, we now add some more discussion on the potential mechanism by which nuclear export inhibition modulates TDP-43 phase separation. This can be found on page 10.

References:

Zhang, Q. et al. A human forebrain organoid model phenocopies dysregulated RNA and protein homeostasis in ALS/FTD-associated TDP-43 proteinopathies. bioRxiv (2025). (https://www.biorxiv.org/content/10.1101/2025.11.09.687455v1.full

Changes from the earlier version to the revised version

Abstract

Earlier version: RNA-binding protein TAR DNA-binding protein 43 (TDP-43) can form liquid-like, nuclear assemblies whose phase transitions are linked to protein aggregation and neurodegeneration. The mechanisms that determine whether these assemblies remain dynamic or convert into solid states are poorly defined. Here we combined chemical and genome-wide genetic screenings to identify cellular processes influencing the phase behavior of an RNA-binding defective TDP-43 mutant associated with amyotrophic lateral sclerosis (ALS) and frontotemporal dementia (FTD). While the screens identified multiple cellular processes—including RNA splicing, translation, and proteostasis—as modulators of TDP-43 phase states, our analyses underscore nuclear transport as a critical determinant of liquid-to-solid transitions. Inhibition of nuclear export prevents cytoplasmic accumulation of immobile TDP-43 assemblies, whereas enhanced export promotes irreversible, solid TDP-43 structures. We validated this mechanism in a brain organoid model of ALS, showing nuclear retention limits pathogenic aggregation and its associated defects. These findings establish nuclear export as a key regulator of TDP-43 phase transitions and define a mechanistic link between altered phase dynamics and neurodegeneration.

Revised version: RNA-binding protein TAR DNA-binding protein 43 (TDP-43) can form liquid-like, nuclear assemblies whose phase behavior may influence its aggregation propensity and neurotoxic activity. The mechanism(s) that modulates the transition of TDP-43 from a liquid to solid phase is poorly defined. Here we combine chemical and genome-wide genetic screenings to identify cellular factors that modulate the phase behavior of an RNA-binding defective TDP-43 mutant that mimics an Amyotrophic Lateral Sclerosis (ALS)-associated variant. Our screens uncover multiple cellular processes including RNA splicing, protein translation, proteostasis imbalance and nuclear export as TDP-43 phase regulators. Importantly, TDP-43 phase transition can be dynamically recapitulated in vitro in a semi-permeabilized cell system, which reveals that the inhibition of nuclear export reshapes the nuclear environment in favor of an RNA-dependent TDP-43 liquid-liquid phase separation (LLPS) state, which mitigates cytoplasmic TDP-43 aggregation. We validated this mechanism in a brain organoid model bearing an ALS-associated mutation, showing that nuclear export deficiency can limit pathogenic phospho-TDP-43 accumulation. These findings establish nuclear export as a key regulator of TDP-43 phase transitions and define a mechanistic framework that links altered nuclear transport and phase dynamics to TDP-43 aggregation potential.

Added in the revised version: The RNA-binding protein TAR DNA-binding protein 43 (TDP-43), encoded by the TARDBP gene, plays a central role in RNA transport, splicing and metabolism 1, 2, and is the major pathological component of cytoplasmic inclusions in amyotrophic lateral sclerosis (ALS) and frontotemporal dementia (FTD) 3–5.

Removed in the revised version: Introduction

Earlier version: The RNA-binding protein TAR DNA-binding protein 43 (TDP-43), encoded by the TARDBP gene, plays a central role in RNA transport, splicing and metabolism (Guo and Shorter, 2017; Suk and Rousseaux, 2020), and is the major pathological component of cytoplasmic inclusions in amyotrophic lateral sclerosis (ALS) and frontotemporal dementia (FTD) patients (Kellett et al., 2025; Neumann et al., 2006; Rummens and Da Cruz, 2025). TDP-43 contains two RNA-recognition motifs (RRMs), a bipartite nuclear-localization sequence, and a glycine-rich C-terminal low-complexity domain (LCD) that is intrinsically prone to aggregation (Tziortzouda et al., 2021). Under physiological conditions, TDP-43 binds RNAs in the nucleus, participating in multiple RNA-metabolic processes. In contrast, disease-causing mutations trigger its mis-localization to the cytoplasm, forming detergent-resistant inclusions. TDP-43-positive protein aggregates were present in ∼97% of ALS and ∼45% of FTD cases, as well as in a subset of Alzheimer’s disease patients. These diseases are now collectively termed TDP-43 proteinopathies (Brettschneider et al., 2014; Neumann et al., 2006). Both loss of nuclear function and gain of cytoplasmic toxicity have been implicated in neurodegeneration (Suk and Rousseaux, 2020; Tziortzouda et al., 2021), underscoring the importance of nucleocytoplasmic transport in regulating TDP-43 patho-physiological functions.

Revised version: TDP-43 contains two RNA-recognition motifs (RRMs), a bipartite nuclear-localization sequence, and a glycine-rich C-terminal low-complexity domain (LCD) that is intrinsically prone to aggregation 6. Under physiological conditions, TDP-43 binds RNAs in the nucleus, participating in multiple RNA-metabolic processes. In contrast, disease-causing mutations trigger its mis-localization to the cytoplasm, forming detergent-resistant inclusions. TDP-43-positive protein aggregates were present in ∼97% of ALS and ∼45% of FTD cases, as well as in a subset of Alzheimer’s disease patients. These diseases are now collectively termed TDP-43 proteinopathies 4, 7. Both loss of nuclear function and gain of cytoplasmic toxicity have been implicated in neurodegeneration 2, 6, underscoring the importance of nucleocytoplasmic transport in regulating TDP-43 patho-physiological functions.

Earlier version: Like many RNA-binding proteins (RBPs), TDP-43 undergoes liquid-liquid phase separation to form membraneless biomolecular condensates (Conicella et al., 2016; Shorter, 2019; Tziortzouda et al., 2021). Under conditions of cellular stress or in the presence of RNA-binding defective mutations, these condensates can transition from dynamic liquid droplets to more rigid gel-like or solid assemblies (Conicella et al., 2016; Gasset-Rosa et al., 2019; Lang et al., 2024). To date, more than 70 disease-linked mutations in TARDBP have been identified in ALS or FTD, many within the LCD, while others clustering near or in the RRMs disrupt RNA binding (Tziortzouda et al., 2021). Defects in RNA binding appears to be a major governor of TDP-43 phase behavior (Babinchak et al., 2019; Mann et al., 2019). Additionally, disease-associated mutations alter TDP-43’s regulatory functions in RNA metabolism, causing widespread cryptic exon inclusion (Ling et al., 2015; Ma et al., 2022; Mehta et al., 2023; Seddighi et al., 2024; Zhang et al., 2025b), alternative polyadenylation and other RNA maturation defects that collectively contribute to neurotoxicity (Zeng et al., 2025).

Revised version: Like many RNA-binding proteins (RBPs), TDP-43 undergoes liquid-liquid phase separation to form membraneless biomolecular condensates 6, 8, 9. Under conditions of cellular stress or in the presence of RNA-binding inhibiting mutations, these condensates can transition from dynamic liquid droplets to rigid gel-like or solid assemblies 8, 10, 11. To date, more than 70 disease-linked mutations in TARDBP have been identified in ALS or FTD, many within the LCD, while others clustering near or in the RRMs to disrupt RNA binding 6. Defects in RNA binding appears to be a major determinant of TDP-43 phase behavior 12, 13. Disease-associated mutations diminish TDP-43’s activities in RNA metabolism, causing widespread cryptic exon inclusion 14–17, alternative polyadenylation and other RNA maturation defects that collectively contribute to neurotoxicity 18, 19.

Earlier version: A major advance in the field of TDP-43 phase regulation is the recent report of a unique form of TDP-43 demixing, forming “anisosomes” that contain an anisotropic spherical shell of an RNA-binding-defective TDP-43 mutant and a central liquid core formed by the HSP70 chaperone (Yu et al., 2021). These structures, observed mostly in the nucleus, were postulated to be an intermediate en route to pathological aggregation when cellular proteostasis or nucleocytoplasmic transport is compromised. Additionally, in vitro studies with purified TDP-43 and selected small molecules have begun to elucidate the molecular basis of TDP-43 self-association (Maurel et al., 2025; Rubien et al., 2022), offering insights into its phase behavior and aggregation mechanisms.

Revised version: Apart from disease-associated mutations, post-translational modifications (PTM), especially acetylation, have been shown to modulate TDP-43’s RNA binding capacity 16, 20. Acetylation was detected on endogenous TDP-43 in ALS patient samples. In vitro, TDP-43 acetylation is enhanced when cells are exposed to oxidative or proteotoxic stress 20, 21. Interestingly, like disease-associated TDP-43 mutants, acetylated TDP-43 is also more aggregation-prone and displays a similar phase regulation pattern 20. For example, a recent study showed that a TDP-43 acetylation mimetic mutant designated as 2KQ and several RNA-binding defective TDP-43 disease variants can all undertake a unique form of demixing, forming “anisosomes” that contain an anisotropic spherical shell of TDP-43 and a central liquid core enriched in the HSP70 chaperone 22. Anisosomes, observed mostly in the nucleus under overexpression conditions, have not been validated in human patient samples. Nevertheless, they might be an intermediate en route to pathological aggregation when cellular proteostasis or nucleocytoplasmic transport is compromised, as suggested by a recent study 22. Additional in vitro studies using purified TDP-43 and small molecules have begun to elucidate the molecular basis of TDP-43 self-association 23, 24, offering insights into its phase behavior and aggregation mechanisms.

Earlier version: Despite these advances, the cellular pathways modulating TDP-43 phase behavior (droplet formation, anisosomal shell–core architecture, conversion to gel/solid structures) and its link to disease-associated aggregation remain poorly defined. Here, we address this knowledge gap by combining a chemical-genetic screen approach with a genome-wide siRNA screen to identify molecular determinants that modulate the phase behavior of an RNA-binding defective TDP-43 mutant. Our study reveals critical contributions of RNA splicing, protein translation, and the HSP90/ubiquitin–proteasome proteostasis network to TDP-43 phase transitions. Importantly, our work demonstrates how nuclear export can influence the transition of TDP-43 from a nuclear liquid-demixed form to a cytoplasmic immobile structure. We validated the impact of nuclear export inhibition on phosphorylated-TDP-43 cytoplasmic accumulation using an iPSC-derived 3-D organoid model bearing an ALS-associated mutation (Zhang et al., 2025b). Collectively, these findings establish a mechanistic framework linking altered phase dynamics and TDP-43 aggregation to nuclear transport defects, a process known to modulate neurodegeneration in ALS and FTD.

Revised version: Despite these advances, the cellular pathways modulating TDP-43 phase behavior (droplet formation, anisosomal shell–core architecture, conversion to gel/solid structures) and its link to disease-associated aggregation remain unclear. We address this knowledge gap by combining a chemical-genetic screen approach with a genome-wide siRNA screen to identify molecular determinants that modulate the phase behavior of the RNA-binding defective TDP-43 2KQ mutant. Our study reveals critical contributions of RNA splicing, protein translation, and the HSP90/ubiquitin–proteasome proteostasis network to TDP-43 phase transition. Importantly, our work demonstrates how nuclear export can influence the transition of TDP-43 from a nuclear liquid-demixed form to a cytoplasmic immobile gel-like structure akin to protein aggregation. Using an iPSC-derived 3-D organoid model bearing an ALS-associated mutation 16, we further validated the impact of nuclear export inhibition on the accumulation of cytoplasmic phosphorylated-TDP-43. Collectively, these findings establish a mechanistic framework linking altered phase dynamics and TDP-43 aggregation to nuclear transport defects, a process known to modulate neurodegeneration in ALS and FTD.

Results

A chemical genetic screen identified modulators of TDP-43 phase behavior

Earlier version: To identify cellular pathways modulating TDP-43 phase behavior, we conducted a chemical genetic screen using a small molecule library targeting diverse known cellular pathways and processes. We employed a previously established DLD1 cell model stably expressing an RNA-binding defective TDP-43 mutant (2KQ) tagged with the green fluorescence protein Clover because TDP-43 2KQ undergoes rapid demixing in the nucleus upon expression induction (Yu et al., 2021). We seeded cells in 384-well plates and induced the expression of TDP-43 2KQ with doxycycline for 24 hours. We then treated these cells with a LOPAC library containing 1280 drugs for 24 hours (Figure 1A). This approach could avoid false positive drugs that affect the induction of TDP-43 2KQ.

Revised version: To identify cellular pathways modulating TDP-43 phase behavior, we conducted a chemical genetic screen using a small molecule library targeting diverse known cellular pathways and processes. We employed a previously established DLD1 cell model stably expressing an RNA-binding defective TDP-43 mutant (2KQ) tagged with the green fluorescence protein Clover. In this system, TDP-43 2KQ undergoes rapid demixing upon induced expression 22, generating 20-30 sphere objects named anisosome in each nucleus. Both the number and size of anisosome increase substantially in the initial phase of TDP-43 induction but remain largely unchanged after 24 h of induction. We seeded cells in 384-well plates and induced the expression of TDP-43 2KQ with doxycycline for 24 hours. We then treated these cells with a LOPAC library containing 1280 drugs for 24 hours (Figure 1A). This approach could avoid false positive drugs that affect the expression of TDP-43 2KQ.

A chemical genetic screen identifies TDP-43 phase modulators

Earlier version: (A) The workflow of the chemical genetic screen. (B) Dose dependent reduction of anisosome number by identified chemicals. DLD1 TDP-43 2KQ-Clover cells were treated with doxycycline to induce anisosome for 24 h and then treated with the drugs as indicated for 15 h. Cells were imaged for anisosome count.

Revised version: (A) Workflow of the chemical genetic screen. (B) Dose dependent reduction of anisosome number by identified chemicals. DLD1 TDP-43 2KQ-Clover cells were treated with doxycycline to induce anisosome for 24 h and then treated with the drugs as indicated for 15 h. Cells were imaged for anisosome count.

Earlier version: High content imaging combined with algorithm-based image analyses (see methods) detected ∼20 spherical anisosomes per cell in controls and most treated cells. However, in cells treated with a subset of the drugs, the number of anisosomes was noticeably reduced (Supplemental Table 1). Importantly, cells treated with Bortezomib, a potent proteasome inhibitor, had fewer but enlarged TDP-43 positive punctae with irregular shapes (Figure 1B), similar to cells treated with another proteasome inhibitor reported previously (van Eersel et al., 2011; Yu et al., 2021). This result validated our screening design.

Revised version: High content imaging combined with algorithm-based image analyses (see methods) detected ∼20 spherical anisosomes per cell in controls and most treated cells. However, in cells treated with a subset of drugs, the number of anisosomes was noticeably reduced (Supplemental Table 1). Importantly, cells treated with Bortezomib, a potent proteasome inhibitor, had fewer but enlarged TDP-43 positive puncta with irregular shapes (Figure 1B) similarly to cells treated with another proteasome inhibitor as reported previously 22, 25. This result validated our screening design.

Earlier version: To further confirm the identified hits, we performed a titration experiment, which showed that these inhibitors dose-dependently reduced the number of anisosomes with IC50 ranging from sub-micromolar to micromolar (Figure 1B, Supplemental Table 1). Interestingly, the screen identified cytoplasmic signaling kinases including PAK4, BTK, and GSK-3β as well as the receptor tyrosine kinase PDGFRβ as modulators of TDP-43 phase behavior (Figure 1B, Supplemental Table 1). PAK4 and GSK-3β are known regulators of the YAP signaling pathway (Jin et al., 2025; Mo et al., 2024), and YAP was recently identified as a TDP-43-interacting protein that influences its phase dynamics (Zhang et al., 2025a). These connections provide additional physiological support for the relevance of the modulators uncovered in our screen. Our screen also linked TDP-43 2KQ demixing to protein translation because cycloheximide, a ribosome elongation inhibitor could modulate TDP-43 phase dynamics (see below).

Revised version: To further confirm the identified hits, we performed a concentration titration experiment, which showed that these inhibitors dose-dependently reduced the anisosome number with IC50 ranging from sub-micromolar to micromolar (Figure 1B, Supplemental Table 1). Immunoblotting showed that treatment with these drugs did not change the total TCP-43 protein level (Supplementary Figure 1), suggesting that the impact on anisosome is not due to altered protein expression. Interestingly, the screen identified cytoplasmic signaling kinases including PAK4, BTK, and GSK-3β as well as the receptor tyrosine kinase PDGFRβ as modulators of TDP-43 phase behavior (Figure 1B, Supplemental Table 1). PAK4 and GSK-3β are known regulators of the YAP signaling pathway 26, and YAP was recently identified as a TDP-43-interacting protein that influences its phase dynamics 27. These connections provide additional physiological support for the relevance of the uncovered modulators. Our screen also linked TDP-43 2KQ demixing to protein translation because cycloheximide, a ribosome elongation inhibitor also affect TDP-43 phase dynamics (see below).

Earlier version: Live cell imaging revealed two distinct types of TDP-43 phase modulators

Revised version: Live cell imaging revealed two distinct classes of TDP-43 phase modulators

Earlier version: We next used confocal microscopy to further characterize the impact of a subset of the identified inhibitors on anisosome dynamics. We chose Tripterin (a drug targeting the heat shock protein HSP90), CP-673451 (a PDGFRβ inhibitor), two deubiquitinase inhibitors VLX-1570 and Spautin-1, and the Exportin-1 inhibitor KPT-276 because previous reports have linked the targets of these compounds to TDP-43 proteinopathies (Archbold et al., 2018; Farrawell et al., 2020; Furukawa et al., 2015; Garcia-Toscano et al., 2024). We first induced anisosome formation for 24 hours and then treated cells with the inhibitors. Compared to DMSO-treated controls, cells treated with Spautin-1 had fewer and smaller anisosomes. KPT-276 treatment also reduced anisosome number, albeit to a lesser extent (Figure 2A, top panels). In contrast, cells treated with CP-673451, Tripterin, or Bortezomib had fewer but larger TDP-43 punctae. Since Tripterin was also reported as a 20S proteasome inhibitor, we confirmed the role of HSP90 in this process by treating cells with the well-established HSP90 inhibitor Geldanamycin, which caused similar phenotypes as Tripterin (Supplemental Figure 1A). These results suggest that TDP-43 phase dynamics are modulated by the cellular folding capacity, the ubiquitin proteasome system, and nucleocytoplasmic transport.

Revised version: We next used confocal microscopy to further characterize a subset of the identified inhibitors. We chose Tripterin (a drug targeting the heat shock protein HSP90), CP-673451 (a PDGFRβ inhibitor), two deubiquitinase inhibitors VLX-1570 and Spautin-1, and the Exportin-1 inhibitor KPT-276 because previous reports have linked the targets of these compounds to TDP-43 proteinopathies 28–31. We first induced anisosome formation for 24 hours and then treated cells with the inhibitors. Compared to DMSO-treated controls, cells treated with Spautin-1 had fewer and smaller anisosomes. KPT-276 treatment also reduced anisosome number, albeit to a lesser extent (Figure 2A, top panels). In contrast, cells treated with CP-673451, Tripterin, or Bortezomib had fewer but larger TDP-43 punctae. Since Tripterin was also reported as a 20S proteasome inhibitor, we confirmed the role of HSP90 in this process using the well-established HSP90 inhibitor Geldanamycin (Supplemental Figure 2A). These results suggest that TDP-43 phase dynamics are modulated by the cellular folding capacity, the ubiquitin proteasome system, and nucleocytoplasmic transport.

Two distinct types of TDP-43 phase modulators.

Earlier version: (A) DLD1 TDP-43 2KQ-Clover cells treated with doxycycline for 24 h were further treated with the indicated compounds for 6 h and imaged (SPA, Spautin-1, 5 μM; KPT, KPT-276, 15 μM; CP, CP-673451, 30 μM; TRP, Tripterin, 1 μM; BTZ, Bortezomib, 10 nM). Scale bar, 5 μm. (B) FRAP analyses of anisosomes after treatment with the indicated drugs for 3-5 h. Circles indicate bleached areas. Scale bars, 1 μm. (C, D) Quantification of the experiments represented in B. FL, Fluorescence. (E) Reverse FRAP analyses of anisosome dynamics. The areas indicated by the dashed lines were bleached. Scale bar, 5 μm. (F) Quantification of the fluorescence loss over time in unbleached anisosomes as shown in E. N=4-6. (G) Quantification of the initial rate of fluorescence loss in unbleached anisosomes. ****, p<0.0001; ***, p<0.001 by unpaired Student’s t-test.

Revised version: (A) DLD1 TDP-43 2KQ-Clover cells treated with doxycycline for 24 h were further treated with the indicated compounds for 6 h and imaged (SPA, Spautin-1, 5 μM; KPT, KPT-276, 15 μM; CP, CP-673451, 30 μM; TRP, Tripterin, 1 μM; BTZ, Bortezomib, 10 nM). Shown are single confocal z-section views. Scale bar, 5 μm. (B) FRAP analyses of anisosomes after treatment with the indicated drugs for 3-5 h. Circles indicate bleached areas. Shown are single confocal z-section views. Scale bars, 1 μm. (C, D) Quantification of the experiments represented in B. MFI, Mean Fluorescence Intensity. N = 5 anisosomes/condition. (E) Reverse FRAP analyses of anisosome dynamics. The areas indicated by the dashed lines were bleached. Shown are single confocal z-section views. Scale bar, 5 μm. (F) Quantification of the Mean Fluorescence Intensity (MFI) loss over time in unbleached anisosomes as shown in E. N = 7 anisosomes/condition. (G) Quantification of the initial rate of fluorescence loss in unbleached anisosomes as shown in E. ****, p<0.0001; ***, p<0.001 by unpaired Student’s t-test (two-tailed). N = 7 anisosomes/condition.

Earlier version: To further dissect the effects of these inhibitors on TDP-43 phase separation, we combined drug treatment with live cell fluorescence imaging, using Fluorescence Recovery After Photobleaching (FRAP) or reversed FRAP to define the dynamics of TDP-43 2KQ. In FRAP, when the green fluorescence of anisosomes was bleached by a laser in untreated cells, it rapidly recovered due to the recruitment of TDP-43 2KQ from the nucleoplasm (Figure 2B, top panels, Figure 2C, D). Conversely, in reversed FRAP, bleaching TDP-43 2KQ surrounding an anisosome and then imaging the anisosome over time revealed a gradual decline in fluorescence intensity within the unbleached anisosome, while the surrounding bleached anisosomes regained some fluorescence (Figure 2E, top panel). This result suggests that TDP-43 2KQ in anisosomes could also be released into the nucleoplasm. Interestingly, in cells treated with Spautin-1 or KPT-276, anisosome-associated TDP-43 2KQ could freely exchange with the nucleoplasmic pool (Figure 2B, D). However, in cells treated with Bortezomib, Tripterin or Geldanamycin, the enlarged TDP-43 puncta appeared static; TDP-43 2KQ fluorescence within anisosomes failed to recover after photobleaching (Figure 2B, C, Supplemental Figure 1B), nor did it decrease over time when surrounding TDP-43 2KQ was bleached (Figure 2E-G). Collectively, these results suggest that these inhibitors affect anisosomes via two distinct mechanisms: one maintains TDP-43 in a demixed liquid state, while the other converts it to a gel-like, solid state.

Revised version: To further dissect the effects of these inhibitors on TDP-43 phase separation, we combined drug treatment with live cell fluorescence imaging, using Fluorescence Recovery After Photobleaching (FRAP) or reversed FRAP. In FRAP, when the green fluorescence of anisosomes was bleached by a laser in untreated cells, it rapidly recovered due to the recruitment of TDP-43 2KQ from the nucleoplasm (Figure 2B, top panels, Figure 2C, D). Conversely, in reversed FRAP, bleaching TDP-43 2KQ surrounding an anisosome repeatedly over time caused a gradual decline in fluorescence intensity within the unbleached anisosome, while the surrounding areas regained fluorescence partially (Figure 2E, top panel). This result suggests that TDP-43 2KQ in anisosomes undergoes rapid exchange with a pool in the nucleoplasm. Likewise, in cells treated with Spautin-1 or KPT-276, anisosome-associated TDP-43 2KQ could freely exchange with the nucleoplasmic pool (Figure 2B, D). In contrast, in cells treated with Bortezomib, Tripterin or Geldanamycin, the enlarged TDP-43 puncta appeared static; TDP-43 2KQ fluorescence within anisosomes failed to recover after photobleaching (Figure 2B, C, Supplemental Figure 2B), nor did it decrease over time when surrounding TDP-43 2KQ was bleached (Figure 2E-G). Collectively, these results suggest that these inhibitors modulate anisosomes via two distinct mechanisms: one maintains TDP-43 in a demixed liquid state, while the other converts it to a gel-like solid state.

A genome-wide siRNA screen identified genetic modulators of TDP-43 phase separation

Earlier version: To further decipher the molecular determinants of TDP-43 phase separation, we conducted an unbiased genome-wide siRNA knockdown (KD) screen (Figure 3A). To this end, we transfected a siRNA library targeting 21,404 human genes each with three siRNAs into DLD1 cells stably expressing Clover-tagged TDP-43 2KQ. After gene KD, anisosomes were induced for 24 hours. High content imaging and automated analysis identified 1,533 candidates whose KD reduced the anisosome number per cell (Z-score > 2) (Supplementary Table S2). Based on a STRING gene network analysis, we selected 211 networked genes with top Z-scores and re-screened each of them with three additional siRNAs. This confirmed the involvement of 110 genes in TDP-43 anisosome regulation (Figure 3A, Supplementary Table S3). GO pathway analyses categorized identified genes into several major pathways including RNA splicing, protein translation, proteasomal degradation, and nuclear transport (Figure 3B). The identification of large number of ribosomal proteins and proteasome subunits is consistent with our chemical genetic screen, which showed that translation and proteasome inhibitors modulate anisosomes (Supplemental Table 1). Interestingly, GO analyses using ‘molecular function’ further linked many identified genes to neurodegenerative diseases, particularly ALS (Figure 3C, D).

Revised version: To further decipher the molecular determinants of TDP-43 phase separation, we conducted an unbiased genome-wide siRNA knockdown (KD) screen (Figure 3A). To this end, we transfected a siRNA library targeting 21,404 human genes each with three siRNAs into DLD1 cells stably expressing Clover-tagged TDP-43 2KQ. After gene KD, anisosomes were induced for 24 hours. High content imaging and automated analysis identified 1,533 candidates whose KD reduced the anisosome number per cell (Z-score > 2) (Supplemental Table S2). To further narrow down the list, we performed a STRING protein network analysis based on the assumption that a protein interaction network bearing multiple positive hits would be more likely to be a true effector. We selected 211 networked genes with top Z-scores and re-screened each of them with three additional siRNAs. This suggested the involvement of 110 genes in TDP-43 anisosome regulation (Figure 3A, Supplementary Table S3). GO pathway analyses categorized these genes into several pathways including RNA splicing, protein translation, proteasomal degradation, and nuclear transport (Figure 3B). The identification of ribosomal proteins and proteasome subunits is consistent with our chemical genetic screen, which implicates translation and proteasomal degradation in anisosome modulation (Supplemental Table 1). Interestingly, GO analyses using ‘molecular function’ linked many identified genes to neurodegenerative diseases, particularly ALS (Figure 3C, D).

A genome-wide siRNA screen identifies modifiers of TDP43 phase behavior.

(A) The workflow of the siRNA genetic screen. GW, genome-wide; AS, anisosome; HT, high-throughput. (B) Pathway analysis of genes whose knockdown reduces anisosome number in cells. The relative anisosome count are indicated by both color and size of the nodes. (C) GO molecular function analysis of TDP-43 phase modifiers. (D) A list of identified genes linked to ALS in C.

Anisosome dynamics is modulated by RNA splicing and protein translation

Earlier version: Given the well-established function of TDP-43 in RNA binding and processing, we investigated the role of RNA splicing in anisosome regulation. To this end, we first incubated Clover-tagged TDP-43 2KQ cells with doxycycline to induce anisosomes and then treated cells with the potent RNA splicing inhibitor Pladienolide-B (PlaB). Confocal microscopy confirmed that RNA splicing inhibition resulted in fewer but larger TDP-43-positive puncta in a dose dependent manner (Figure 4A, B, Supplemental Figure 2A). Despite the size difference, anisosomes in drug-treated cells were morphologically indistinguishable from controls. FRAP experiments further showed that TDP-43 within PlaB-treated condensates remained highly mobile, although the fluorescence recovery rate was slightly reduced compared to controls (Figure 4C, D). Time-lapse live cell imaging frequently detected fusion of TDP-43 puncta after PlaB treatment (Figure 4E, Movie S1). Together, these results suggest that disrupting RNA splicing stabilizes TDP-43 in demixed liquid state, resulting in larger condensates.

Revised version: Given the well-established function of TDP-43 in RNA binding and processing, we investigated the role of RNA splicing in anisosome regulation. To this end, we first incubated TDP-43 2KQ cells with doxycycline to induce anisosomes and then treated cells with a potent RNA splicing inhibitor Pladienolide-B (PlaB). Confocal microscopy confirmed that RNA splicing inhibition resulted in fewer but larger TDP-43-positive puncta in a dose dependent manner (Figure 4A, B, Supplemental Figure 3A).

Anisosome phase behavior is modulated by RNA splicing and protein translation/

Earlier version: (A) The splicing inhibitor Pladienolide-B (PlaB) reduces anisosome number in a dose dependent manner. Representative images of anisosome-induced cells treated with DMSO (control), 5 nM, or 20 nM PlaB for 16 h. Scale bar: 10 µm. (B) Quantification showing the number of anisosome (AS) per cell in TDP-43 2KQ-Clover cells treated with PlaB as indicated. * p <0.05, ** p <0.01, **** p < 0.0001 by One-way ANOVA. n=3 biological repeats. (C) Representative FRAP images of anisosomes in cells treated for 16 h with DMSO or PlaB (20 nM). BB, before photobleaching, right after photobleaching (0 s), or 4 and 30 seconds after photobleaching (4 s and 30 s). Scale bar, 1 µm. (D) The graph shows the quantification of the remaining TDP-43 fluorescence (FL) in C. Error bars indicate mean ± SD, N = 28 for control and 23 for PlaB-treated cells. (E) Live cell imaging of anisosome fusion in TDP-43 2KQ-Clover cells treated with 20 nM PlaB for 5 h before tracking the fusion. Representative images from Movie S1 showing fusion events indicated by arrows. Scale bar, 1 µm. (F) Representative images of anisosome-induced (24 h) cells treated with DMSO (control) or ANS (200 nM) for 16 h. Scale bar, 1 µm. (G) Quantification of the number of anisosomes per cell in randomly selected images of DMSO-, Cycloheximide (CHX)-, or ANS-treated cells. **** p < 0.0001 by one-way ANOVA; each dot represents an image with at least 20 cells counted. n=3 biological repeats. (H) Quantification of anisosome size in control or cells treated with CHX or ANS. ****, p < 0.0001 by One-way ANOVA. n=3 biological repeats. AU, arbitrary unit. (I) As in C except that anisosome-induced cells were treated with CHX or ANS before photobleaching. Scale bar 1 µm. (J) Quantification of fluorescence recovery in CHX- or ANS-treated cells vs the DMSO control.

Revised version: (A) The splicing inhibitor Pladienolide-B (PlaB) reduces anisosome number in a dose dependent manner. Representative images of anisosome-induced cells treated with DMSO (control), 5 nM, or 20 nM PlaB for 16 h. Shown are maximum intensity projection views. Scale bar: 10 µm. (B) Quantification showing the number of anisosome (AS) per cell in TDP-43 2KQ-Clover cells treated with PlaB as indicated. * p <0.05, ** p <0.01, **** p < 0.0001 by ordinary one-way ANOVA and Dunnett’s multiple comparisons. N=3 biological repeats, and in each experiment, at least 30 randomly selected cells were analyzed. (C) Representative z-section confocal FRAP images of anisosomes in cells treated for 16 h with DMSO or PlaB (20 nM). BB, before photobleaching, right after photobleaching (0 s), or 4 and 30 seconds after photobleaching (4 s and 30 s). Scale bar, 1 µm. (D) The graph shows the quantification of the remaining TDP-43 fluorescence (FL) in C. Error bars indicate mean ±SD, N = 28 for control and 23 for PlaB-treated cells. MFI, Mean Fluorescence Instensity. (E) Live cell imaging of anisosome fusion in TDP-43 2KQ-Clover cells treated with 20 nM PlaB for 5 h before tracking the fusion. Representative reconstructed 3-D images from Movie S1 showing fusion events indicated by arrows. Scale bar, 1 µm. (F) Representative maximum intensity projection views of anisosome-induced (24 h) cells treated with DMSO (control) or ANS (200 nM) for 16 h. Scale bar, 1 µm. (G) Quantification of the number of anisosomes per cell in randomly selected images of DMSO-, Cycloheximide (CHX)-, or ANS-treated cells. **** p < 0.0001 by ordinary one-way ANOVA and Dunnett’s multiple comparisons. Each dot represents a randomly selected field with at least 20 cells counted from one of the 3 biological repeats. (H) Quantification of anisosome size in control or cells treated with CHX or ANS. ****, p < 0.0001 by ordinary one-way ANOVA and Dunnett’s multiple comparisons. N=3 biological repeats. AU, arbitrary unit. (I) As in C except that anisosome-induced cells were treated with CHX or ANS before photobleaching. Scale bar 1 µm. (J) Quantification of fluorescence recovery in CHX- or ANS-treated cells vs the DMSO control. N=10 anisosomes/condition.

Added in the revised version: Despite size difference, anisosomes in drug-treated cells were morphologically indistinguishable from controls. FRAP experiments further showed that TDP-43 within PlaB-treated condensates remained highly mobile, although the fluorescence recovery rate was slightly reduced compared to controls (Figure 4C, D). Time-lapse live cell imaging frequently detected fusion of TDP-43 puncta after PlaB treatment (Figure 4E, Movie S1). Together, these results suggest that disrupting RNA splicing stabilizes TDP-43 in a demixed liquid state, resulting in larger condensates.

Earlier version: Next, we explored the effect of translation inhibition on anisosome dynamics using Cycloheximide (CHX) and Anisomycin (ANS). Confocal imaging showed that after anisosome induction, translational shutdown by CHX or ANS did not significantly reduce TDP-43 expression, suggesting a long half-life. Similar to PlaB-treated cells, CHX- and ANS-treated cells contained fewer, enlarged anisosomes (Figure 4F-H, Supplemental Figure S2B). FRAP analysis revealed that TDP-43 dynamics within anisosomes were similar between ANS-treated and control cells (Figure 4I, J). These results suggest that translation inhibition also stabilizes TDP-43 in liquid condensates.

Revised version: Next, we explored the effect of translation inhibition on anisosome dynamics using Cycloheximide (CHX) and Anisomycin (ANS), two translation elongation inhibitors. Confocal imaging and immunoblotting showed that translation inhibition by CHX or ANS did not significantly reduce TDP-43 expression (Supplemental Figure 1, 3). Like PlaB-treated cells, CHX- and ANS-treated cells contained fewer, enlarged anisosomes (Figure 4F-H, Supplemental Figure S2B). FRAP analysis revealed that TDP-43 dynamics within anisosomes were similar between ANS-treated and control cells (Figure 4I, J). These results suggest that translation inhibition also stabilizes TDP-43 in liquid condensates.

Earlier version: Nuclear export regulates TDP-43 liquid-to-solid phase transition

Revised version: Nuclear export modulates TDP-43 liquid-to-solid phase transition

Our screen also identified several nuclear transport regulators (e.g., Exportin-1/XPO1, Exportin-2/CSE1L) and nuclear pore components as anisosome regulators (Figure 3B). We focused on XPO1 because our chemical genetic screen identified two XPO1 inhibitors, KPT-276 and Verdinexor as potent anisosome modulators (Figure 1B).

Earlier version: To explore the role of XPO1 in TDP-43 anisosome regulation, we induced anisosome formation in Clover-TDP-43 2KQ cells and treated them with Leptomycin B (LMB), a potent XPO1 inhibitor. LMB treatment resulted in a time- and dose-dependent reduction in anisosome number, while increasing their sizes (Figure 5A-C). The enlarged TDP-43 puncta showed the typical hollowed anisosome ring structure (Figure 5D), indicating that TDP-43 remained in the liquid state. Indeed, FRAP experiments demonstrated that LMB treatment did not significantly affect the fluorescence recovery rate of TDP-43 within bleached anisosomes (Figure 5D, E). Time-lapse confocal imaging showed that shortly after LMB treatment, preformed anisosomes began to fuse with each other (Figure 5F, Movie S2). Thus, XPO1 inhibition stabilizes TDP-43 in a liquid phase.

Revised version: To explore the role of XPO1 in TDP-43 anisosome regulation, we induced anisosome formation in Clover-TDP-43 2KQ cells and treated them with Leptomycin B (LMB), a potent XPO1 inhibitor. LMB treatment resulted in a time-and dose-dependent reduction in anisosome number, while increasing their sizes (Figure 5A-C). The enlarged TDP-43 puncta showed the typical hollowed anisosome ring structure (Figure 5D), indicating that TDP-43 remained in liquid state. Indeed, FRAP experiments demonstrated that LMB treatment did not significantly affect the fluorescence recovery rate of TDP-43 within bleached anisosomes (Figure 5D, E). Time-lapse confocal imaging showed that shortly after LMB treatment, preformed anisosomes began to fuse with each other (Figure 5F, Movie S2). Thus, XPO1 inhibition stabilizes TDP-43 in a liquid phase.

XPO1 regulates anisosome liquid-to-solid transition.

Earlier version: (A) Pharmacological inhibition of XPO-1 with Leptomycin B (LMB) reduces the number of anisosome. The graph on top indicates the experimental design. Arrows show cells with enlarged anisosomes. Scale bar 10 µm (B) Quantification of anisosome numbers per cell in experiments represented by A. Error bars indicate mean ± SD; **** p <0.0001 by one-way ANOVA. n=3 biological repeats. (C) LMB dose dependently reduces anisosome number. Anisosome were induced in TDP-43 2KQ-Clover cells followed by treatment with LMB at the indicated concentrations for 16 h. The histogram shows the distribution of cells (50-90 cells/condition) by anisosome count. (D) FRAP experiments demonstrate that anisosomes remain in a liquid phase following LMB treatment (200 nM, 16 h). Anisosome-induced (24 h) cells were treated with DMSO or LMB for 16 h and then photobleached at the indicated areas. Scale bar 1 µm. (E) Quantification of the FRAP experiment in D. (F) Time-lapse confocal microscopy detects anisosome fusion after LMB (200 nM, 5 h) treatment in TDP-43 2KQ-Clover cells. Arrows indicate fusion events. Scale bar, 1 µm (G) Overexpression of mCherry-tagged XPO-1 in TDP-43 2KQ-Clover cells induces cytoplasmic TDP-43 puncta. TDP-43 2KQ-Clover cells (green) transfected with mCherry-XPO1 (red) were stained with DAPI (blue) to label nuclei (dashed lines). Cells were imaged 48 h post-transfection. Panels 1, 2 show a representative confocal section, while panels 3-6 show reconstructed 3-D views. The position and volume of anisosomes were also presented in magenta in a surface-rendered view (SRV) in panel 5. The arrow in panel 3 indicates an example of cytoplasmic TDP-43 aggregate. Scale bar, 10 µm. (H) Quantification of the percentage of cells showing cytosolic TDP-43 puncta in XPO-1 positive (pos) and negative (neg) cells in randomly selected fields from 3 independent experiments. **** p <0.0001 by unpaired Student’s t-test. (I) FRAP-based confocal imaging reveals the transition of anisosomes into a gel-like state upon XPO-1 overexpression. Scale bar, 1 µm. (J) Anisosome formation changes endogenous XPO-1 localization. TDP-43 2KQ-Clover before or after anisosome induction were stained with anti-XPO-1 antibodies (red) and DAPI (blue). The bleached areas were indicated by dashed lines. Scale bar, 10 µm. (K) Quantification of nuclear endogenous XPO-1 levels in individual cells (indicated by dots) before or after anisosome induction. **** p < 0.001 by unpaired Student’s t-test. n=2 independent repeats. (L) A schematic model depicting how the nuclear XPO-1 activity influences TDP-43 liquid-to-phase transition. AS, anisosome.

Revised version: (A) Pharmacological inhibition of XPO-1 with Leptomycin B (LMB, 200 nM) reduces the number of anisosome. The graph on top indicates the experimental design. Arrows show cells with enlarged anisosomes. Scale bar 10 µm (B) Quantification of anisosome numbers per cell in experiments represented by A. Error bars indicate mean ±SD; **** p <0.0001 by ordinary one-way ANOVA and Dunnett’s multiple comparisons. N=3 biological repeats. (C) LMB dose dependently reduces anisosome number. Anisosome were induced in TDP-43 2KQ-Clover cells followed by treatment with LMB at the indicated concentrations for 16 h. The histogram shows the distribution of cells (50-90 cells/condition) by anisosome count. (D) FRAP experiments demonstrate that anisosomes remain in a liquid phase following LMB treatment (200 nM, 16 h). Anisosome-induced (24 h) cells were treated with DMSO or LMB for 16 h and then photobleached at the indicated areas. Scale bar 1 µm. (E) Quantification of the FRAP experiment in D. N=18 for control and 16 anisosomes for LMB-treated cells. (F) Time-lapse confocal microscopy detects anisosome fusion after LMB (200 nM, 5 h) treatment in TDP-43 2KQ-Clover cells. Arrows indicate fusion events. Scale bar, 1 µm (G) Overexpression of mCherry-tagged XPO-1 in TDP-43 2KQ-Clover cells induces cytoplasmic TDP-43 puncta. TDP-43 2KQ-Clover cells (green) transfected with mCherry-XPO1 (red) were stained with DAPI (blue) to label nuclei (dashed lines). Cells were imaged 48 h post-transfection. Panels 1, 2 show a representative confocal section, while panels 3-6 show reconstructed 3-D views. The position and volume of anisosomes were also presented in magenta in a surface-rendered view (SRV) in panel 5. The arrow in panel 3 indicates an example of cytoplasmic TDP-43 aggregate. Scale bar, 10 µm. (H) Quantification of the percentage of cells showing cytosolic TDP-43 puncta in XPO-1 positive (pos) and negative (neg) cells in randomly selected fields from 3 independent experiments. **** p <0.0001 by two-tailed unpaired Student’s t-test. (I) FRAP-based confocal imaging reveals the transition of anisosomes into a gel-like state upon XPO-1 overexpression. Scale bar, 1 µm. (J) Anisosome formation changes endogenous XPO-1 localization. TDP-43 2KQ-Clover before or after anisosome induction were stained with anti-XPO-1 antibodies (red) and DAPI (blue). The bleached areas were indicated by dashed lines. Scale bar, 10 µm. (K) Quantification of nuclear endogenous XPO-1 Mean Fluorescence Intensity (MPI) in individual cells (indicated by dots) before or after anisosome induction. **** p < 0.001 by two-tailed unpaired Student’s t-test. N = 58 cells for uninduced and 63 cells for induced condition in 2 independent repeats. (L) A schematic model depicting how the nuclear XPO-1 activity influences TDP-43 liquid-to-phase transition. AS, anisosome.

Earlier version: To corroborate the inhibitor study, we overexpressed mCherry-tagged XPO1. Surprisingly, we observed fewer and enlarged TDP-43 positive puncta in mCherry-XPO1-expressing cells (Figure 5G, panels 1, 2). At first glance, this phenotype resembled that of XPO1-inhibited cells. However, careful examination revealed several distinct features. First, despite the size increase, we did not observe the typical hollowed anisosome ring structure. Secondly, while anisosomes in mCherry-XPO1 negative cells were almost exclusively nuclear, ∼30% of mCherry-XPO1-expressing cells contained cytoplasm-localized TDP-43 puncta (Figure 5G, panels 1, 2, 3, Figure 5H). Thirdly, FRAP experiments demonstrated that Clover-TDP-43 fluorescence did not recover after photobleaching, suggesting that TDP-43 was in a gel-like state transitioning to aggregates (Figure 5I). Notably, mCherry-XPO1 could be detected within TDP-43 puncta in some cells (Figure 5G, panels 3-6). These findings are consistent with a previous report (Archbold et al., 2018), demonstrating that XPO1 promotes TDP-43 nuclear egression and cytoplasmic aggregate formation. Immunostaining of endogenous XPO1 in anisosome-induced cells further confirmed this conclusion. Although antibodies could not penetrate anisosomes (Yu et al., 2021), anisosome induction depleted XPO1 from the nucleoplasm, presumably due to its sequestration into anisosomes (Figure 5J, K). Since XPO1 does not bind TDP-43 directly (Pinarbasi et al., 2018), additional factors likely facilitate XPO1-mediated TDP-43 nuclear egression under this condition. Collectively, these results suggest that the stability and dynamics of anisosomes are modulated by XPO1: reduced XPO1 activity stabilizes the liquid state and favors large TDP-43 condensates, whereas increasing XPO1 activity promotes its transition to a gel-like structure and cytoplasmic redistribution (Figure 5L). This conclusion underscores a critical regulatory link between XPO1-dependent nuclear egression and the TDP-43 liquid-to-solid phase transition.

Revised version: To corroborate the inhibitor study, we overexpressed mCherry-tagged XPO1. Surprisingly, we observed fewer and enlarged TDP-43 positive puncta in mCherry-XPO1-expressing cells (Figure 5G, panels 1, 2). At first glance, this phenotype appeared similar to that of XPO1-inhibited cells. However, careful examination revealed several distinct features. First, despite the size increase, we did not observe the typical hollowed anisosome ring structure. Instead, these large puncta appeared more irregular in shape. Secondly, while anisosomes in mCherry-XPO1 negative cells were almost exclusively nuclear, ∼30% of mCherry-XPO1-expressing cells contained cytoplasm-localized TDP-43 puncta (Figure 5G, panels 1, 2, 3, Figure 5H). Thirdly, FRAP experiments demonstrated that Clover-TDP-43 fluorescence did not recover after photobleaching, suggesting that TDP-43 was in a gel-like state in transition to aggregates (Figure 5I). Notably, mCherry-XPO1 could be detected within TDP-43 puncta in some cells (Figure 5G, panels 3-6). These findings are consistent with a previous report that showed the regulation of TDP-43 nuclear egress by multiple nuclear export receptors including XPO1 31. As expected, over-expression of mCherry did not alter the subcellular localization of TDP-43 puncta, nor did mCherry co-localize with TDP-43 (Supplemental Figure S4).

Added in the revised version: Immunostaining of endogenous XPO1 in anisosome-induced cells showed that anisosome induction depleted XPO1 from the nucleoplasm (Figure 5J, K), further suggesting a possible connection between XPO-1 and TDP-43. Since XPO1 does not bind TDP-43 directly 32, an indirect interaction between these molecules might cause the depletion of XPO-1 from nucleoplasm in anisosome-induced cells. However, since antibody staining could not conclusively demonstrate the sequestration of endogenous XPO-1 in anisosomes due to an antibody penetration barrier 22, alternative interpretation cannot be excluded. Collectively, our results suggest that the stability and dynamics of anisosomes are modulated by XPO1-mediated nuclear export: reduced XPO1 activity stabilizes a liquid state and favors larger TDP-43 condensates, whereas increasing XPO1 activity promotes its transition to a gel-like structure (Figure 5J).

Added in the revised version: Nuclear export inhibition stabilizes TDP-43 anisosomes in an RNA dependent manner

Added in the revised version: To further dissect the role of nuclear export in TDP-43 phase regulation, we developed a semi-permeabilized cell–based in vitro assay (Figure 6A). To this end, anisosomes were first induced in Clover-tagged TDP-43 2KQ cells, followed by selective permeabilization of the plasma membrane using the pore-forming toxin streptolysin O (SLO). Permeabilization was monitored at 37 °C by confocal microscopy using the membrane-impermeable dye NucSpot, which stained nuclei only upon plasma membrane disruption.

Added in the revised version: The inhibition of XPO-1 stabilizes anisosomes in an RNA-dependent manner.

Added in the revised version: (A) A schematic diagram of the semi-permeabilized cell assay. SLO, streptolysin O. (B) TDP-43 in anisosomes becomes soluble after cell permeabilization (indicated by NucSpot positive staining in magenta). Note: The intensity of NucSpot dye indicates permeabilization time. Dashed circles highlight cells permeabilized early during the incubation. The cells have fewer anisosomes. Scale bars, 5μm. (C) TDP-43 mean fluorescence intensity (MPI) displays an inverse correlation with the NucSpot dye intensity in B. AU, arbitrary unit. (D) Small TDP-43 positive puncta were reformed after permeabilized DLD1 cells were incubated with cow liver cytosol in the presence of an ATP regeneration system and GTP (100 μM) (ATP/GTP). Cells were fixed and stained with Hoechst (blue) to label nuclei. Scale bars, 5 μm. (E) Quantification of the TDP-43 Mean Fluorescence Intensity (MPI) in cells as shown in D. Error bars, SD. ****, p<0.0001; ***, p<0.001 by unpaired Student’s t-test. N=at least 35 randomly selected cells representing two independent experiments. (F) LMB treatment (200 nM, 16 h) stabilizes anisosomes. Dashed circles indicate permeabilized cells. Scale bars, 5 μm. (G) RNAase T1 treatment destabilizes anisosomes in DKD1 cells pre-treated with LMB (200 nM, 16 h). DLD1 cells induced for TDP-43 expression were treated with LMB (200nM, 16 h). Cells were permeabilized in the absence or presence of RNAase T1 for 40 min before imaging. Images shown in this figure are maximum intensity projection views of confocal z-sections. Scale bars, 5 μm.

Added in the revised version: Notably, permeabilization led to a marked reduction in TDP-43–positive nuclear puncta and overall fluorescence intensity (Figure 6B). Quantitative analysis revealed an inverse correlation between NucSpot signal and TDP-43 fluorescence (Figure 6C). The reduction in TDP-43 signal was not due to protein elimination. Instead, TDP-43 appeared to shift from a phase-separated state to a soluble state, likely because of cytosolic factor loss and ATP depletion during cell permeabilization. Consistent with this interpretation, supplementation of permeabilized cells with cow liver cytosol together with an ATP-regenerating system (ARS) and GTP restored TDP-43 puncta (Figure 6D), albeit much smaller than anisosomes. Nevertheless, this led to a significantly increase in fluorescence intensity (Figure 6E).

Added in the revised version: The re-formed puncta may represent intermediates en route to anisosome assembly. However, these structures failed to fuse into larger anisosomes even after prolonged incubation, suggesting that cytosol add-back only partially restores the intracellular environment required for full anisosome maturation. Addition of buffer supplemented with ARS and GTP also modestly induced the formation of small TDP-43 puncta and increased fluorescence intensity (Figure 6D, E). In contrast, buffer or cytosol alone was insufficient to support puncta formation. Together, these results establish this semi-permeabilized system as a platform in which TDP-43 condensates can be dissolved and partially reconstituted under controlled conditions.

Added in the revised version: We next applied this assay to examine the mechanism underlying leptomycin B (LMB)-induced enlargement of anisosomes. Strikingly, anisosomes in LMB-treated cells remained stable following permeabilization (Figure 6F). However, inclusion of RNase T1, which degrades single-stranded RNA, during permeabilization led to their dissolution (Figure 6G). These findings suggest that LMB stabilizes anisosomes, at least in part, by increasing nuclear RNA availability.

Earlier version: XPO1 inhibition mitigates TDP-43 hyperphosphorylation and modest rescues neuronal defects in an anisosome-bearing organoid model

Revised version: Nuclear export inhibition mitigates TDP-43 hyperphosphorylation in an organoid model of TDP-43 proteinopathy

Earlier version: Since increased XPO1 activity is associated with more cytoplasmic, gel-like TDP-43, we tested whether reducing XPO1 activity could mitigate TDP-43 cytosolic aggregation using a recently established anisosome-bearing organoid model of TDP-43 proteinopathies. This model is derived from an induced pluripotent stem cell (iPSC) line carrying the ALS-associated mutation (K181E) (Zhang et al., 2025b), which like 2KQ disrupts the RNA binding activity of TDP-43 (Chen et al., 2019). Additionally, the K181E mutant accumulates in a liquid-condensed phase in association with HSP70, and like anisosomes, these condensates can convert to solid aggregates during cellular stress or HSP70 depletion (Gu et al., 2021).

Revised version: Since increasing the XPO1 activity causes more cytoplasmic, gel-like TDP-43, we postulate that reducing XPO1 activity might mitigate TDP-43 cytosolic aggregation. To test this idea, we used a recently established anisosome-bearing organoid model of TDP-43 proteinopathies 16, which is derived from an induced pluripotent stem cell (iPSC) line carrying the ALS-associated mutation (K181E) in the endogenous TDP-43 locus 16. Like the 2KQ mutant, the K181E substitution disrupts the RNA binding activity of TDP-43 33. Additionally, the K181E mutant was reported to accumulate in a liquid-condensed phase in association with HSP70 in cultured cells under overexpression conditions, and like anisosomes 22, these condensates can be converted to solid aggregates during cellular stress or HSP70 depletion 34.

Earlier version: We treated homozygous K181E organoids or wild-type (WT) controls with KPT-276 at a low dose (20 nM) for 35 days, starting on day 87 based on pilot data indicating good tolerance with prolonged treatment. Since hyperphosphorylated TDP-43 (p-TDP-43) is a major aggregate component, we stained organoids with antibodies against p-TDP-43 and total TDP-43. Confocal imaging revealed that WT and K181E organoids had similar amounts of total TDP-43, but p-TDP-43 was only significantly present in mutant organoids (Figure 6A), as reported previously (Zhang et al., 2025b). KPT-276 did not alter the total TDP-43 levels in either genotype (Figure 6A) but significantly reduced the cytosolic p-TDP-43 positive puncta in K181E organoids (Figure 6A, B).

Revised version: We treated organoids bearing the endogenous K181E mutation or wild-type (WT) controls with KPT-276 at a low dose (20 nM) for 35 days starting on day 87 based on pilot data indicating good tolerance with prolonged treatment. We then sectioned and stained these organoids with antibodies against Serine 409/410-phosphorylated TDP-43 (p-TDP-43) and total TDP-43 since hyperphosphorylated TDP-43 (S409/410) is a clinically validated pathological marker for TDP-43-containing cytoplasmic aggregates. In this experiment, homozygous organoids were chosen to enhance the detection of endogenous p-TDP-43. Confocal imaging revealed that WT and K181E organoids had similar amounts of total TDP-43, but p-TDP-43 was only significantly present in mutant organoids (Figure 7A), as reported 16. KPT-276 did not alter the total TDP-43 levels in either genotype (Figure 7A) but significantly reduced the p-TDP-43 positive puncta in K181E organoids (Figure 7A, B). This result suggests that nuclear export is required for the formation of p-TDP-43-containing aggregates in a disease relevant organoid model.

Earlier version: Inhibition of XPO-1 reduces TDP-43 hyperphosphorylation but only partially restore neuronal gene expression in K181E organoids

Revised version: Inhibition of XPO-1 mitigates TDP-43 hyperphosphorylation in TDP-43 K181E organoids.

Earlier version: (A) Confocal fluorescence imaging reveals significant reduction in phosphorylated TDP-43 (p-TDP43) in K181E/K181E organoids following KPT-276 treatment. Organoids of the indicated genotypes at day 87 were treated with DMSO as a control, or with KPT-276 (20 nM) for 35 days. Organoids were fixed, sectioned, and stained with antibodies against TUJ1, a neuronal marker (Magenta), TDP-43 (green) and p-TDP-43 (red). n= 3-5 organoids from two individual batches. (B) Quantification of p-TDP-43 Mean Spot Volume (top) and percentage of cells lacking nuclear TDP-43 (bottom) in experiments represented by A. Error bars indicate mean ± s.e.m., * p < 0.05, **** p < 0.0001 by one-way ANOVA, n=3 organoids. (C) Gene Ontology (GO) analysis of genes whose expression was disrupted in K181E/K181E organoids but restored by KPT-267 as shown in Supplemental Fig. S3. BP, Biological Pathway; CC, Cellular Component; MF, Molecular Function.

Revised version: (A) Confocal fluorescence imaging reveals significant reduction in phosphorylated TDP-43 (p-TDP43) in K181E/K181E organoids following KPT-276 treatment. Organoids of the indicated genotypes at day 87 were treated with DMSO as a control, or with KPT-276 (20 nM) for 35 days. Organoids were fixed, sectioned, and stained with antibodies against TUJ1, a neuronal marker (Magenta), TDP-43 (green) and p-TDP-43 (red). n= 3-5 organoids from two individual batches. (B) Quantification of p-TDP-43 Mean Spot Volume (top) and percentage of cells lacking nuclear TDP-43 (bottom) in experiments represented by A. Error bars indicate mean ±s.e.m. * p < 0.05, **** p < 0.0001 by one-way ANOVA, n=3 organoids/each.

Removed in the revised version: To assess whether KPT-276 treatment also rescued the neuronal defects in TDP-43 K181E organoids, we performed bulk RNAseq using KPT-276-treated K181E and WT organoids. A comparison of treated mutant organoids to WT or untreated K181E organoids showed that KPT-276 reversed expression defects in 58 downregulated (2.90%) and 45 (1.99 %) upregulated genes (Supplemental Figure 3, Supplemental Table S4). GO pathway analysis of genes rescued from downregulation by KPT-276 highlighted neuronal pathways including neuron differentiation and brain development; for genes upregulated in K181E organoids, rescued genes are linked to G-protein coupled neurotransmitter receptor activities at the plasma membranes (Figure 6C, Supplementary Table S4). Together, these results demonstrate that XPO1 inhibition by KPT-276 modestly restores gene transcription defects in a subset of neuronal pathways in K181E mutant organoids, while robustly mitigating FTD-associated p-TDP-43 accumulation. Thus, reducing TDP-43 cytoplasmic aggregation alone may offers only limited benefit in disease treatment.

Discussion

Earlier version: ALS- and FTD-associated TDP-43 mutants defective in RNA binding are known to undergo self-association through their intrinsically disordered regions and RNA-binding domains, generating demixed liquid, gel, and aggregated states that remain in dynamic exchange. This protein phase behavior is shared by many RNA binding proteins and is increasingly recognized as a contributor to protein aggregation in neurodegenerative diseases (Shorter, 2019). Intriguingly, TDP-43 phase separation forms unique layered nuclear structures, with an HSP70-filled central core surrounded by a protein shell made of TDP-43. Although it is unclear how these structures are nucleated in the cell, our observations suggest that, once formed, they can grow by recruiting additional TDP-43 from the surrounding environment. Small anisosomes may also fuse with one another, although such fusion events become infrequent and difficult to detect after anisosomes reach a steady state,

Revised version: ALS- and FTD-associated TDP-43 mutants defective in RNA binding are known to undergo self-association through their intrinsically disordered regions and RNA-binding domains, generating demixed liquid, gel, and aggregated states that remain in dynamic exchange. This protein phase behavior is shared by many RNA binding proteins and is increasingly recognized as a contributor to protein aggregation in neurodegenerative diseases 9. Intriguingly, TDP-43 phase separation forms unique layered nuclear structures, with an HSP70-filled central core surrounded by a protein shell made of TDP-43. Although it is unclear how these structures are nucleated in the cell, our observations suggest that, once formed, they can grow by recruiting additional TDP-43 from the surrounding environment. Small anisosomes may also fuse with one another, although such fusion events become infrequent and difficult to detect after anisosomes reach a steady state.

Earlier version: Our study supports the notion that TDP-43, at high concentrations, can transition from a demixing-prone, liquid state into a gel-like state en route to forming amorphous aggregates (Babinchak et al., 2019; Carey and Guo, 2022). We found that inhibition of RNA splicing, protein translation, or nuclear export favors the liquid state of TDP-43, while proteasome and HSP90 inhibition or enhancing XPO-1-mediated nuclear egression promotes the conversion of anisosomes to a gel-like state.

Revised version: Our study supports the notion that TDP-43, at high concentrations, can change from a demixing-prone, liquid state into a gel-like state en route to forming amorphous aggregates 12, 35. We found that proteasome and HSP90 inhibition or enhancing XPO-1-mediated nuclear egression promotes the conversion of TDP-43 anisosomes into a gel-like state. By contrast, inhibition of RNA splicing, protein translation, or nuclear export all favors the liquid state of TDP-43.

Added in the revised version: The roles of RNA splicing and nuclear export in TDP-43 phase regulation might be mechanistically linked because anisosomes stabilized upon nuclear export inhibition are sensitive to RNAase treatment in semi-permeabilized cells. Anisosome maturation probably involves one or more RNA splicing intermediate. Therefore, the completion of RNA splicing followed by nuclear egression of spliced RNAs would disfavor anisosome growth and maturation.

Earlier version: Because the TDP-43 2KQ mutant lacks RNA binding activity, the effects of splicing inhibition are likely mediated by other RNA-binding factors involved in splicing regulation. TDP-43 is known to interact with many splicing factors including members of the hnRNP family like hnRNP A1/A2 and PSF (Freibaum et al., 2010); Inhibition of RNA splicing may alter TDP-43’s interactions with these factors, trapping it in certain complexes that favor its retention in the liquid state. Alternatively, splicing inhibition may change the overall nuclear environment, allowing large anisosomes to fuse with each other as observed in our study. This could explain why PlaB treatment reduces the number of TDP-43 anisosomes while maintaining their dynamic liquid property. Likewise, inhibition of nuclear export may influence anisosome dynamics by reshaping the nuclear environment, lowing the barrier to anisosome fusion. Future studies are required to elucidate the precise mechanisms by which these pathways regulate anisosome dynamics.

Revised version: Because the TDP-43 2KQ mutant lacks RNA binding activity, anisosome maturation is likely modulated by other RNA-binding factors. In this regard, TDP-43 is known to interact with many splicing factors including members of the hnRNP family like hnRNP A1/A2 and PSF 36; Inhibition of RNA splicing might alter TDP-43’s interactions with these factors, trapping it in an RNA-containing complex that favors anisosome growth and maturation, This could explain why both PlaB and LMB reduce the number of TDP-43 anisosomes while maintaining their dynamic liquid property. Future studies are required to elucidate the precise mechanism by which nuclear transport and RNA splicing modulate anisosome dynamics.

Earlier version: Our validation of proteasome inhibition as a driver of TDP-43 droplet enlargement and hardening reinforces the idea that impaired protein clearance promotes the transition from liquid condensates to gel-like or solid aggregates. This is consistent with previous studies in ALS models, which showed that proteasome stress accelerates TDP-43 aggregation and toxicity (van Eersel et al., 2011). In contrast, inhibition of deubiquitinases reduced the number of TDP-43 anisosomes, probably by promoting the proteasome-mediated clearance of TDP-43. The role of HSP90 in this process may be linked to TDP-43 folding stability, as HSP90 can bind TDP-43 to prevent promiscuous interactions (Lin et al., 2021). Collectively, these observations suggest that TDP-43 condensates are maintained by a delicate balance between formation and clearance, modulated by proteolytic and chaperone activities.

Revised version: Our validation of proteasome inhibition as a driver of TDP-43 droplet enlargement and hardening reinforces the idea that impaired protein clearance promotes the transition from liquid condensates to gel-like or solid aggregates. This is consistent with previous studies in ALS models, which showed that proteasome stress accelerates TDP-43 aggregation and toxicity 25. In contrast, inhibition of deubiquitinases reduced the number of TDP-43 anisosomes, probably by promoting the proteasome-mediated clearance of TDP-43. The role of HSP90 in this process may be linked to TDP-43 folding stability, as HSP90 can bind TDP-43 to prevent promiscuous interactions 37. Collectively, these observations suggest that TDP-43 condensates are maintained by a delicate balance between formation and clearance, modulated by proteolytic and chaperone activities.

Earlier version: While out chemical and genetic screens have identified multiple pathways influencing TDP-43 phase states, a particularly compelling aspect of our study is the discovery that the nuclear export receptor XPO1 governs TDP-43 liquid-to-solid transitions and subcellular localization. Pharmacological inhibition of XPO1 mitigates cytoplasmic p-TDP-43 accumulation, whereas XPO1 overexpression promotes the cytoplasmic accumulation of TDP-43-containing puncta. However, XPO1 may modulate TDP-43 nuclear egression indirectly, as suggested previously (Duan et al., 2022; Pinarbasi et al., 2018). While previous studies have not detected strong interactions between XPO-1 and TDP-43 (Pinarbasi et al., 2018), a weak or indirect interaction between XPO-1 and TDP-43 may still exist, especially when TDP-43 undergoes demixing. This model would explain why anisosome induction depletes endogenous XPO-1 from the nucleoplasm. Our findings also hint at a potential feedback mechanism in which TDP-43 demixing perturbs nuclear export, analogous to how cytoplasmic TDP-43 aggregates disrupt nuclear import. Given that ALS is characterized by cytoplasmic TDP-43 mis-localization and hyperphosphorylation, the XPO-1-dependent nuclear export pathway may constitute a key molecular switch linking physiological phase behavior to pathological aggregation.

Revised version: While out chemical and genetic screens have identified multiple pathways influencing TDP-43 phase states, a particularly compelling aspect of our study is the discovery that the nuclear export receptor XPO1 governs TDP-43 liquid-to-solid transitions and subcellular localization. Pharmacological inhibition of XPO1 mitigates cytoplasmic p-TDP-43 accumulation, whereas XPO1 overexpression promotes the cytoplasmic accumulation of TDP-43-containing puncta. However, XPO1 may modulate TDP-43 nuclear egression indirectly, as suggested previously 32, 38. While previous studies have not detected strong interactions between XPO-1 and TDP-43 32, a weak or indirect interaction between XPO-1 and TDP-43 may still exist, especially when TDP-43 undergoes demixing. This model would explain why anisosome induction depletes endogenous XPO-1 from the nucleoplasm. Our findings also hint at a potential feedback mechanism in which TDP-43 demixing perturbs nuclear export, analogous to how cytoplasmic TDP-43 aggregates disrupt nuclear import. Given that ALS is characterized by cytoplasmic TDP-43 mis-localization and hyperphosphorylation, the XPO-1-dependent nuclear export pathway may constitute a key molecular switch linking physiological phase behavior to pathological aggregation.

Earlier version: Notably, in 3-D brain organoids bearing the K181E mutation, inhibition of nuclear export by KPT-276 only modestly rescues the transcriptional defects in homozygous TDP-43 K181E organoids despite an almost complete ablation of cytoplasmic phospho-TDP-43. This observation suggests that cytoplasmic TDP-43 aggregation is only a modest contributor to TDP-43 proteinopathies that lead to neurodegeneration. Major defects may still be attributed to deregulated RNA processing due to altered RNA binding (Hergesheimer et al., 2019; Joseph et al., 2025; Mehta et al., 2023; Zhang et al., 2025b). Thus, effective therapy likely requires both inhibition of cytoplasmic TDP-43 aggregation and restoration of normal RNA processing and nuclear transport. Together, our results support a model in which multiple cellular pathways including proteostasis, RNA metabolism, translation, and nuclear export—cooperatively determine the phase behavior of TDP-43. These findings have deepened our understanding of how physiological phase separation can evolve into pathogenic aggregation through cumulative failures of cellular quality control and nuclear export systems.

Revised version: Together, our results support a model in which multiple cellular pathways including proteostasis, RNA metabolism, protein translation, and nuclear export—cooperatively determine the phase behavior of TDP-43. These findings have deepened our understanding of how physiological phase separation can evolve into pathogenic aggregation through cumulative failures of cellular quality control and nuclear export systems.

Methods

RNAi-based genetic screen for anisosome regulators

Earlier version: Genome-wide RNA interference (RNAi) screen was performed using DLD1 cells stably expressing Clover-tagged TDP-43 2KQ. Cells were transfected with siRNAs targeting 21,404 genes for 72 hours, followed by induction of anisosome formation with doxycycline (0.5 µg/mL) for 24 h. Specifically, the screen was carried out in 384-well format with 3 individual siRNAs for each gene. To prepare siRNA transfection, control siRNA (40 nmol) was mixed with 2 mL water to create a 50X (20 µM) stock and siRNA-death (20 nmol) was mixed with 1 mL of water for the same concentration. Dilute siRNA stocks to a 400 nM working concentration in water. For each well, mix 0.1 µL transfection reagent in 20 µL of cell culture medium without FBS or P/S. Add 2 µL of siRNA solution to each well and incubate with RNAiMAX for 30 minutes at room temperature. While incubating RNAiMAX, prepare trypsinized cells by ensuring thorough separation and accurate cell counting. Add 20 µL of cell suspension at a density of 0.43 × 105 cells/mL (850 cells/well) in medium containing 20% FBS and 2 x P/S to achieve a final volume of 40 µL per well. Incubate cells for 3 days before addition of 10 µL 10 x doxycycline solution to each well (final volume 50 µL) for anisosome induction. 24 h later, high-throughput (HT) imaging and automated phenotypic screening were conducted to identify genes that modulated anisosome formation.

Revised version: Genome-wide RNA interference (RNAi) screen was performed using DLD1 cells stably expressing Clover-tagged TDP-43 2KQ. Cells were transfected with siRNAs targeting 21,404 genes for 72 hours, followed by induction of anisosome formation with doxycycline (0.5 µg/mL) for 24 h. Specifically, the screen was carried out in 384-well format with 3 individual siRNAs for each gene. To prepare siRNA transfection, control siRNA (40 nmol) was mixed with 2 mL water to create a 50X (20 µM) stock and siRNA-death (20 nmol) was mixed with 1 mL of water for the same concentration. Dilute siRNA stocks to a 400 nM working concentration in water. For each well, mix 0.1 µL transfection reagent in 20 µL of cell culture medium without FBS or P/S. Add 2 µL of siRNA solution to each well and incubate with RNAiMAX for 30 minutes at room temperature. While incubating RNAiMAX, prepare trypsinized cells by ensuring thorough separation and accurate cell counting. Add 20 µL of cell suspension at a density of 0.43 ×105 cells/mL (850 cells/well) in medium containing 20% FBS and 2 x P/S to achieve a final volume of 40 µL per well. Incubate cells for 3 days before addition of 10 µL 10 x doxycycline solution to each well (final volume 50 µL) for anisosome induction. 24 h later, high-throughput (HT) imaging and automated phenotypic screening were conducted to identify genes that modulated anisosome formation. To control for siRNA transfection efficiency, we routinely included a positive knockdown control in which cells were treated with a mixture of siRNAs targeting several essential genes (AllStars Hs Cell Death siRNA (Qiagen, #1027299).

Fluorescence imaging was conducted using an Opera Phenix High-Content Screening System (PerkinElmer). Whole-well images were acquired in confocal mode using a 20X water objective lens after anisosome induction for 24 h. Images were reconstructed and analyzed using the PerkinElmer Columbus server (v2.9.1). Nuclei were segmented based on Hoechst33342 staining, and anisosomes were detected via YFP signals within the nuclear region. The mean anisosome count per nucleus was calculated to assess changes in anisosome levels, while toxicity was evaluated by counting the total number of nuclei. Data were normalized to the siRNA-Neg control for each plate, and Z-scores were calculated to identify significant hits. Identified 1,533 genes whose knockdown significantly reduced the number of anisosomes per cell (Z-score > 2) were subjected to protein network analysis by STRING to identify candidate anisosome regulatory genes that were subject to a second-round screen.

A chemical genetic screen for anisosome regulators

We also performed a chemical genetic screen using LOPACR1280 (Sigma) compound library on DLD1 cells expressing Clover-tagged TDP-43-2KQ. Briefly, cells were induced with doxycycline for 24 h to promote anisosome formation, followed by treatment with individual compounds for an additional 24 hours. High-throughput (HT) imaging and automated phenotypic screening were conducted to identify compounds that modulated anisosome formation.

Visualization and analysis of anisosome dynamics

To examine the impact of drug treatment on anisosome size, cells were induced to form anisosome for 24 h and then treated by different inhibitors as follows: cycloheximide (CHX, 20 µg/mL), anisomycin (ANS, 200 nM) for 16 hours. For splicing inhibition, cells were exposed to Pladienolide-B (PlaB) at concentrations of 5 or 20 nM, or DMSO (control) for 16 hours before imaging. Cells were fixed with 4% paraformaldehyde in Phosphate Buffer Saline (PBS). Randomly selected fields were used to quantify the number of anisosomes per cell with at least two biological repeats.

To image anisosome fusion, live cell imaging was performed by treating cells with Leptomycin B (200 nM) or PlaB (20 nM) for 5 hours and then imaged by confocal 3-D sectioning for 30 minutes to capture anisosome fusion events. For Fluorescence Recovery After Photobleaching (FRAP), cells were treated with DMSO, Leptomycin B, Geldanamycin or PlaB as indicated in figure legends. Either part of an anisosome or entire anisosome was photobleached and imaged. TDP-43 fluorescence recovery after photobleaching was quantified by ImageJ and calculated to assess anisosome dynamics. At least 10 anisosomes for each condition were bleached and analyzed.

To evaluate the time course of Leptomycin B (LMB)-induced anisosome changes, DLD1 TDP-43 2KQ-Clover cells were seeded and induced with 0.5 µg/mL doxycycline for 24 hours. Cells were treated with 10 µM LMB and fixed at 0-, 8-, and 16-hours post-treatment. Fixed cells were stained with Hoechst to visualize nuclei. 10-20 confocal 3-D sections were obtained for each randomly selected field. Images were converted to maximum projected view by ImageJ before anisosome counting and intensity measurement. The density plot of anisosome counts per cell was generated using R, illustrating dose-dependent effects of leptomycin on anisosome formation.

To test the effect of XPO1 overexpression on TDP-43 phase regulation, DLD1 TDP-43 2KQ-Clover cells were co-transfected with mCherry-tagged XPO1 (mCh-XPO1) for 24 h before doxycycline was added at 0.5 µg/mL to induce anisosome formation for 24 h before confocal imaging.

To study the relationship between anisosome induction and endogenous XPO1 localization, cells were seeded, induced with 0.5 µM doxycycline for 48 hours, and fixed with paraformaldehyde in PBS. Immunostaining was performed using an XPO1 antibody (Cell Signaling, 46249S, 1:125) and Hoechst nuclear stain. Fluorescence intensity of XPO1 in both induced and uninduced cells was measured by ImageJ.

Added in the revised version: Semi-permeabilized cell experiments

Added in the revised version: To reconstitute anisome disassembly and re-assembly in seminar permeabilized cells, 30,000 DLD1 cells were seeded in an eight-well ibidi cell chamber two days before the experiment. Cells were treated with Doxycycline for 24 h to induce anisosomes and then treated with inhibitors as indicated in the figure legends. We then washed Cells three times with ice-cold PBS containing 2 mM MgCl2 (PBS-Mg) and then treated these cells in the same buffer (300 µl) containing 200 units of SLO (Sigma-Aldrich, cat. no. SAE0089) on ice for 30 min. The cells were then washed two times with PBS-Mg buffer and incubated with 200 µl of reagent mixture containing 200 µl of PBS in the absence of presence of ARS/GTP or 200 µl of cow liver cytosol (CLC) in the presence of the NucSpot dye. Cells were either imaged directly or fixed with 4% paraformaldehyde before imaging by a Nikon CSU-W1 SoRa super-resolution confocal microscope.

Experiments with organoids

Earlier version: Stem cell gene editing of the patient K181E mutation in iPSC cells was reported previously (Zhang et al., 2025b). Differentiation of iPSCs into forebrain neurons was performed following a previously published protocol (Fernandopulle et al., 2018). Briefly, the procedure used an induction medium (IM-N2) and a neuronal culture medium (CM). IM-N2 was prepared with Knockout DMEM/F12, N2 supplement, NEAA, Gluta-MAX, Chroman I, and doxycycline. On Day 0, iPSCs were observed for confluency, dissociated with Accutase, centrifuged, and resuspended in IM-N2 with Chroman I, and then seeded into Matrigel-coated plates. Over the next four days, cells were monitored microscopically for neurite extensions while media containing doxycycline is refreshed daily. By Day 4, cells exhibit neurite growth and are ready for replating onto poly-L-ornithine (PLO)-coated dishes, prepared in advance by coating with PLO solution, incubating, washing, and drying. Replated cells were cultured in CM comprising BrainPhys medium, B27+ supplement, neurotrophic factors (GDNF, BDNF, NT-3), laminin, and doxycycline, to promote neuronal maturation.

Revised version: Stem cell gene editing of the patient K181E mutation in iPSC cells was reported previously 16. Differentiation of iPSCs into forebrain neurons was performed following a previously published protocol39. Briefly, the procedure used an induction medium (IM-N2) and a neuronal culture medium (CM). IM-N2 was prepared with Knockout DMEM/F12, N2 supplement, NEAA, Gluta-MAX, Chroman I, and doxycycline. On Day 0, iPSCs were observed for confluency, dissociated with Accutase, centrifuged, and resuspended in IM-N2 with Chroman I, and then seeded into Matrigel-coated plates. Over the next four days, cells were monitored microscopically for neurite extensions while media containing doxycycline is refreshed daily. By Day 4, cells exhibit neurite growth and are ready for replating onto poly-L-ornithine (PLO)-coated dishes, prepared in advance by coating with PLO solution, incubating, washing, and drying. Replated cells were cultured in CM comprising BrainPhys medium, B27+ supplement, neurotrophic factors (GDNF, BDNF, NT-3), laminin, and doxycycline, to promote neuronal maturation.

Earlier version: 3-D culture and organoid growth were performed as previously described (Nguyen, 2022). In brief, iPSCs grown on Matrigel were dissociated into single cell suspension by Versene solution (ThermoFisher) and seeded into a 12-well Aggrewell plate (Stemcell Technologies) at 4,000 cells/ well. Next day, spheroids were transferred to an ultralow attachment plate (Corning) containing phase I medium: DMEM/F12, 20% Gibco KnockOut Serum Replacement, 1X Glutamax, 1X MEM Nonessential Amino Acid, 55 µM β-Mercaptoethanol, 1X Pen/Strep, 2 µM Dorsomorphine, 2 µM A83-01. After 5 days, medium was switched to phase II medium: DMEM/F12, 1X N-2 Supplement, 1X Glutamax, 1X MEM Nonessential Amino Acid solution, 1X Pen/Strep, 4 ng/mL WNT3a, 1 µM CHIR-99021, 1 µM SB-431542. On day 7, spheroids were embedded into Matrigel and allowed to continue growing for 7 more days. On day 14, individual spheroids were manually freed from the Matrigel and transferred to a SpinOmega bioreactor spinning at 120 RMP with phase III media: DMEM/F12, 1X N-2 Supplement, 1X B-27 Supplement, 1X Glutamax, 1X MEM Nonessential Amino Acid, 55 µM 2-Mercaptoethanol, 1X Pen/Strep, 2.5 µg/mL insulin. 50 days later, the medium was switched to final differentiation medium: Neurobasal medium, 1X B-27 Supplement, 1X Glutamax, 55 µM 2-Mercaptoethanol, 1X Pen/Strep, 0.2 mM Ascorbic acid, 0.5 mM cAMP, 20 ng/mL brain-derived neurotrophic factor (BDNF), 20 ng/mL glial-derived neurotrophic factor (GDNF). 107-day-old organoids were dissociated into single cells using a 50:50 mixed of Accutase and 0.25% Trypsin with DNase I (1 mg/mL) and plated onto chambered glass slides pretreated with 1% Matrigel (Corning, Inc). Alternatively, organoids were fixed and sectioned for immunostaining (see below). For drug treatment, forebrain organoids (87 Day) were transferred to a 12-well plate on an orbital shaker (120 RPM) and treated with 20 nM KPT276 in final differentiation medium. Medium was changed every other day (Dexoregen, Inc). The organoids were harvested and frozen after 35 days of treatment for sectioning or RNAseq analysis.

Revised version: 3-D culture and organoid growth were performed as previously described 40. In brief, iPSCs grown on Matrigel were dissociated into single cell suspension by Versene solution (ThermoFisher) and seeded into a 12-well Aggrewell plate (Stemcell Technologies) at 4,000 cells/ well. Next day, spheroids were transferred to an ultralow attachment plate (Corning) containing phase I medium: DMEM/F12, 20% Gibco KnockOut Serum Replacement, 1X Glutamax, 1X MEM Nonessential Amino Acid, 55 µM β-Mercaptoethanol, 1X Pen/Strep, 2 µM Dorsomorphine, 2 µM A83-01. After 5 days, medium was switched to phase II medium: DMEM/F12, 1X N-2 Supplement, 1X Glutamax, 1X MEM Nonessential Amino Acid solution, 1X Pen/Strep, 4 ng/mL WNT3a, 1 µM CHIR-99021, 1 µM SB-431542. On day 7, spheroids were embedded into Matrigel and allowed to continue growing for 7 more days. On day 14, individual spheroids were manually freed from the Matrigel and transferred to a SpinOmega bioreactor spinning at 120 RMP with phase III media: DMEM/F12, 1X N-2 Supplement, 1X B-27 Supplement, 1X Glutamax, 1X MEM Nonessential Amino Acid, 55 µM 2-Mercaptoethanol, 1X Pen/Strep, 2.5 µg/mL insulin. 50 days later, the medium was switched to final differentiation medium: Neurobasal medium, 1X B-27 Supplement, 1X Glutamax, 55 µM 2-Mercaptoethanol, 1X Pen/Strep, 0.2 mM Ascorbic acid, 0.5 mM cAMP, 20 ng/mL brain-derived neurotrophic factor (BDNF), 20 ng/mL glial-derived neurotrophic factor (GDNF). 107-day-old organoids were dissociated into single cells using a 50:50 mixed of Accutase and 0.25% Trypsin with DNase I (1 mg/mL) and plated onto chambered glass slides pretreated with 1% Matrigel (Corning, Inc). Alternatively, organoids were fixed and sectioned for immunostaining (see below). For drug treatment, forebrain organoids (87 Day) were transferred to a 12-well plate on an orbital shaker (120 RPM) and treated with 20 nM KPT276 in final differentiation medium. Medium was changed every other day (Dexoregen, Inc). The organoids were harvested and frozen after 35 days of treatment for sectioning and immunostaining. Each experiment was performed with two batches of organoids per condition to reduce the effect of batch variation. The wildtype cells and K181E mutant have the same genetic background.

Tissue preparation and immunostaining

Organoids were fixed in 4% paraformaldehyde in PBS for 1 hour at room temperature, washed with PBS, and immersed in 15% sucrose/PBS solution overnight. Subsequently, organoids were embedded in O.C.T. compound (Sakura) in a plastic mold and frozen down in an ultralow freezer. Embedded organoids were sliced onto glass slides using a Cryostat (Leica). Slides were rinsed with PBS, permeabilized with 0.5% Triton-X/PBS solution for 1 hour at room temperature and blocked using 1% donkey serum in 0.1% Tween-20/PBS solution for 1 hour. Primary antibodies, chicken anti-TUJ1 (Abcam, Ab41489, 1:1,000 dilution), rabbit anti-TDP-43 (10782-2-AP, 1/1000 dilution), mouse anti-phospho-TDP-43 (Cosmobio, CAC-TIP-PTD-M01A, 1/500 dilution), were added to the slides and the slides were incubated in a humidified chamber at 4 °C overnight. After several washes in 0.1% Tween-20/PBS solution, DAPI (Sigma) or secondary antibodies, donkey anti-rabbit, anti-mouse, and anti-chicken (Jackson ImmunoResearch), diluted 1:1,000 in blocking solution, were added to the slides. After 1 hour of incubation at room temperature, the slides were washed and mounted in antifade mounting solution (Fisher scientific).

Image acquisition and processing

Earlier version: Fluorescence confocal images were acquired using a Nikon CSU-W1 SoRa microscope equipped with temperature and CO2 control enclosures and a 60x TIRF lens. Image reconstructions and analyses, including 3-D and time lapse visualizations, were performed using Imaris software or Image J. Fluorescence intensity was quantified using the open-source Fiji/Image J software. Images of randomly selected fields were split into individual channels. A consistent threshold method was applied to individual channel. The particle analysis function was used to automatically identify all anisosomes or other fluorescence structures for size and intensity measurement. Statistical analyses were conducted using Excel (for Student’s t-test) or GraphPad Prism versions 8.0 and 9.0. P-values were calculated using Student’s t-test in Excel or one-way ANOVA in GraphPad Prism. Curve fitting, including linear and nonlinear models, as well as IC50 calculations, were also performed using GraphPad Prism versions 8.0 and 9.0.

Revised version: Fluorescence confocal z-section images were acquired using a Nikon CSU-W1 SoRa microscope equipped with temperature and CO2 control enclosures and a 60x TIRF lens. Randomly selected fields were imaged each with ∼25 slices. Maximum intensity projection views were reconstructed and analyzed by the open-source Image J software (also named Fiji). 3-D and time lapse visualizations were achieved by the Imaris software. Fluorescence intensity was quantified using Fiji. To automatedly count anisosomes, maximum intensity projection views were split into individual channels by Fiji. A consistent default threshold method was applied to each channel. The particle analysis function was used to automatically identify anisosomes or other fluorescence structures for size and intensity measurement.

Added in the revised version: Statistical analyses were conducted using Excel (for Student’s t-test) or GraphPad Prism versions 8.0, 9.0, 11. P-values were calculated using Student’s t-test in Excel or one-way ANOVA in GraphPad Prism. Curve fitting, including linear and nonlinear models, as well as IC50 calculations, were also performed using GraphPad Prism versions 8.0, 9.0 and 11.

Removed in the revised version: Bulk RNA-seq library preparation and data analysis

Removed in the revised version: 800 ng of total RNA samples to generate the library with the NEBNext rRNA Depletion Kit v2 (Human/ mouse/rat) (NEB #E7405), NEBNext Ultra II Directional RNA Library Prep Kit for Illumina (NEB #E7760) and NEBNext® Multiplex Oligos for Illumina® (96 Unique Dual Index Primer Pairs, NEB, E6440) as described in the manufacturer’s manual. Sequencing carried out in the Illumina NextSeq2000 instrument with 101×101 pair end configuration by NCI genomic core.

Removed in the revised version: Adapters were trimmed using Cutadapt (v4.7), and reads were aligned to the GRCh38 genome build with STAR (v2.7.11b) using gene annotations from GENCODE v42. Gene expression was quantified with HTSeq (v2.0.7) (Putri et al., 2022), utilizing the same gene annotations. For differential gene expression analysis, all samples were processed uniformly following the standard DESeq2 (Love et al., 2014) workflow. Differential expression was defined based on an adjusted P-value threshold of <0.05 and |log2 fold change (log2FC) > 1. To identify gene expression rescued by exportin-1 inhibitor treatment, the fold change of expression level should be less than without inhibitor treatment, with P-value < 0.05.

Statistic and reproducibility

All statistical analyses were conducted with GraphPad Prism v10 or Excel. Statistical methods and the number of cells or brain organoid samples (N) are indicated in figure legends or shown in figures as individual data point. Biological repeats (n) are specified in figure legends. No biological repeat was excluded from the analyses. All experiments were repeated at least twice with individual data point labeled in figures unless specified. For imaging analyses, cells in randomly selected fields were analyzed. The researchers were not blinded. The iPSC cell line was derived from a ADRDs genetic risk-free clone, which was initially obtained from a male. Figures were prepared using ImageJ 1.54f, Imaris 9.9.0, Adobe Photoshop v25.12.1, and Adobe Illustrator 28.7.4.